<a id="home"></a>
## Assignment 1
In this assignment you will need to implement functions throughout the classification flow.<br/>

To receive YOUR specific task, go to section 1c - [Your personal implementation tasks](#1c), 
<b>after</b> you run the preceding step - [import packages](#Preceding) and implement & run the [myName](#myName), [myId](#myId) functions. <br/>

<b>Index of functions you need to implement (total possible points: 7)</b>:<br/>

| Section | Section-name | points | Section | Section-name | points | Section | Section-name | points |
| :- | :- | :- | :- | :- | :- | :- | :- | :- | 
| Guidelines | [General assignment instructions](#gen_instructions) | | Imports | [Import packages](#Preceding) | | Part 1. | [Student info](#student_info) | 0.5 |
| 1c. | [info: YOUR personal implementation tasks](#1c) | | Part 2. | [Load the dataset](#load_dataset) | 0.5 | Part 3. | [Scaling](#scaling) | 1.5 |
| Part 4.a. | [Calculate Test,Train Distance](#dist) | 1.5 | Part 4.b. | [KNN main flow - predict](#predict) | 2 | Part 5. | [Model Evaluation](#evaluation) | 1 |

<a id="gen_instructions"></a>
## Working on the assignment, todo and not todo:
<b><u>Todo</u></b>:<br />
<div class="alert alert-success">
<ul><li>&#x2714; <b> Work by yourself</b> and submit your own assignment, <b>no pairing</b> to other students</li>
<li>&#x2714; Test and save your assignment - <b>submit the last tested and saved version</b></li>
<li>&#x2714; Work on the <b>original notebook</b> of the assignment</li>
<li>&#x2714; submit <b>only</b> the <b>notebook (.ipynb)</b> file</li></ul>
</div>
<b><u>Not todo</u></b>:<br />
<div class="alert alert-danger">
<ul><li>&#x2718; do <b>NOT</b> submit an <b>empty assignment</b></li>
<li>&#x2718; do <b>NOT submit extra files</b>, <u>unless you're asked to do so</u></li>
<li>&#x2718; <b>Do NOT submit a .py/.txt/.rar/.zip (or any non (.ipynb) file)</b> <u>version for the notebook</u> of the assignment</li>
<li>&#x2718; <b>do NOT change</b> the notebook file name</li></ul>
</div>

<b>Not good - Any of the following</b>:<br/>
<div class="alert alert-danger">
<ul><li>&#x2718; Code not Implemented</li>
<li>&#x2718; No returned value</li>
<li>&#x2718; Did not remove the '<code>raise NotImplementedError()</code>' line</li></ul>
<code># --------------------------- (AFTER YOUR IMPLEMENTATION) RUN THIS CODE CELL  -------------------------------------
'''
What do you need to do?

Complete the 'method_to_implement' function to do something (some explanation ...) and
     return that value back
'''
def method_to_implement():
    # YOUR CODE HERE
    raise NotImplementedError()
</code></div>

<b>Good - All of the following</b>:<br/>
<div class="alert alert-success">
<ul>
<li>&#x2714; Your code starts one line after the '<code># YOUR CODE HERE</code>' remark</li>
<li>&#x2714; Removed the '<code>raise NotImplementedError()</code>' line</li>
<li>&#x2714; Implemented code</li>
<li>&#x2714; returned expected value</li></ul>
<code># --------------------------- (AFTER YOUR IMPLEMENTATION) RUN THIS CODE CELL  -------------------------------------
'''
What do you need to do?

Complete the 'method_to_implement' function to do something (some explanation ...) and
     return that value back
'''
def method_to_implement():
    # YOUR CODE HERE
    implemented code ...
</code></div>

[Go to the beginning of the assignment](#home)
## Preceding Step - import modules (packages)
This step is necessary in order to use external modules (packages). <br/>
**No other modules are allowed**

**Use the following libraries for the assignment, when needed**:

<b>Import packaged - permitted packages</b>:<br />
<div class="alert alert-success">
&#x2714; Any package or module imported in the following code cell are permitted
</div>
<hr>
<b>Import packaged - Illegal packages</b>:<br/>
<div class="alert alert-danger">
&#x2718; Any package or module NOT IMPORTED in the following code cell are ILLEGAL
</div>

```python
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# === CELL TYPE: IMPORTS AND SETUP 

import sys                       # for testing use only
import os                        # for testing use only
from datetime import datetime    # for testing use only
import random                    # for testing use only
import hashlib                   # for testing use only
import pandas as pd
import numpy as np
import math 
import statistics
```

[Go to the beginning of the assignment](#home)
<a id="student_info"></a>
<h3>Part 1 - Student information methods and YOUR personal implementation:</h3><br/>
The following 2 cells consist of 2 methods which aim to validate your details.<br/>
Please <b>make sure to implement <u>both</u> of them</b><br/><br/>
<table align='left'>
    <tr><td>
        <img src="./images/id-card.png" alt="name and id" width="150" align='left'/>
    </td></tr>
</table>    

### Part 1.a. Instructions
<a id="myName"></a>
<u>Part 1.a.</u> Student information methods; 
<u>method name</u>: <b>myName</b><br/>
<pre>
Complete the following function to return your full name (as a string)
    For example: assume your name is John Smith, you should write:
    return 'John Smith'
------------
return value:
- your full name (as a string)
---------------------
</pre>

<!-- grade_id: answer-myName, points: 0, grade: false, solution: true -->
```python
# 1.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ANSWER 

def myName():
    # YOUR CODE HERE
    return 'Roi Susi'
```

```python
# 1.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: SELF TESTING
# Use the following code to test your implementation:
student_name = myName()
# --- add additional code to check your code if needed:
# YOUR CODE HERE

```

### Part 1.b. Instructions
<a id="myId"></a>
<u>Part 1.b.</u> Student information methods; 
<u>method name</u>: <b>myId</b><br/>
<pre>
Complete the following function to return your ID number (as an integer number)
    For example: assume your ID number is 1234, you should write:
    return 1234
------------
return value:
- your ID number (as an integer number)
---------------------
</pre>

<!-- grade_id: answer-myId, points: 0, grade: false, solution: true -->
```python
# 1.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ANSWER 

def myId():
    # YOUR CODE HERE
    return 300685906
```

```python
# 1.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: SELF TESTING
# Use the following code to test your implementation:
student_id = myId()
# --- add additional code to check your code if needed:
# YOUR CODE HERE

```

<!-- grade_id: test1-1_student_info, points: 0.5, grade: true, solution: false -->
```python
# 1.a. 1.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 1.a. 1.b. - Test 1 (name: test1-1_student_info, points: 0.5)")
print ("\t--->Testing the implementation of 'myName' and 'myId' ...")

# dataframe for output:
dt1 = datetime.now()
try:
    student_name, student_id = myName(), myId()
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"

s_datetime =  [dt1.strftime('%Y-%m-%d %H:%M:%S'), datetime.now().strftime('%Y-%m-%d %H:%M:%S')]
df_output = pd.DataFrame({'value': [student_name, student_id], 'type': [type(student_name), type(student_id)], 'date_time': s_datetime},index=['student_name','student_id'])

print ("\nGood Job!\nYou've passed the test for the implementation of 'myName' and 'myId'  :-)")

print ('\nOutput dataframe:')
print ('-----------------')
df_output

```

Output:
```
Part 1.a. 1.b. - Test 1 (name: test1-1_student_info, points: 0.5)
	--->Testing the implementation of 'myName' and 'myId' ...

Good Job!
You've passed the test for the implementation of 'myName' and 'myId'  :-)

Output dataframe:
-----------------
                  value           type            date_time
student_name   Roi Susi  <class 'str'>  2021-04-11 11:49:27
student_id    300685906  <class 'int'>  2021-04-11 11:49:27
```

[Go to the beginning of the assignment](#home)
<a id="1c"></a>
## 1.c. Info: Your personal implementation tasks - <u>implement the following methods</u> ...
Run the following code cell to and you will know what do YOU NEED to implement

```python
# 1.c.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# === CELL TYPE: IMPLEMENTATION TASKS

# --------------------------- RUN THIS CODE CELL (IMPLEMENTATION INFO) -------------------------------------
# Run the following to get information about your implementation
# ---------------------------
def get_assignment_params(student_id):
    num_modulo = 2 ** 31-1
    created_1st_state = int(hashlib.md5(str(student_id).encode('utf-8')).hexdigest(),16) % num_modulo
    created_2nd_state = int(hashlib.md5(str(created_1st_state).encode('utf-8')).hexdigest(),16) % num_modulo
    created_3rd_state = int(hashlib.md5(str(created_2nd_state).encode('utf-8')).hexdigest(),16) % num_modulo
    scale_types, dist_methods, eval_metrics, binary_cols, score_cols = ['t-distribution standardization', 'minmax normalization'], ['manhattan', 'euclidean', 'chebyshev'], ['accuracy', 'error_rate', 'precision', 'recall'], np.array(['gender_num', 'lunch_type', 'has_preparations']), np.array(['math_score', 'reading_score', 'writing_score'])
    
    assignment_params = {}
    assignment_params['num_train'] = 800
    assignment_params['index_col'] = 'student_id'
    assignment_params['o_features'] = ['social_grp', 'parent_edu']
    
    np.random.seed(created_1st_state)
    np.random.seed(123)
    assignment_params['indices_1st'] = list(np.random.choice(12000, 1000, replace=False))
    np.random.seed(created_2nd_state)
    np.random.seed(345)
    assignment_params['indices_2nd'] = list(np.random.permutation(1000))
    
    random.seed(created_3rd_state)
    random.seed(456)
    y_col = random.choice(binary_cols)
    assignment_params['y_col'] = y_col
    assignment_params['binary_cols'] = list(binary_cols[binary_cols!=y_col])
    
    np.random.seed(created_1st_state)
    np.random.seed(789)
    indx_score = sorted(np.random.choice(3, 2, replace=False))
    selected_score_cols = score_cols[indx_score]
    assignment_params['score_col_1'] = selected_score_cols[0]
    assignment_params['score_col_2'] = selected_score_cols[1]
    random.seed(created_1st_state)
    random.seed(789)
    assignment_params['if_correlated_filter'] = random.choice(selected_score_cols)
    
    random.seed(created_2nd_state)  
    assignment_params['scale_type'] = random.choice(scale_types)
    
    random.seed(created_3rd_state)  
    assignment_params['dist_method'] = random.choice(['manhattan', 'euclidean'])
    
    random.seed(created_1st_state)  
    assignment_params['eval_metric'] = random.choice(eval_metrics)
    
    cols_no_filter = assignment_params['o_features'] + assignment_params['binary_cols'] + list(selected_score_cols)
    cols_w_filter = assignment_params['o_features'] + assignment_params['binary_cols'] + [assignment_params['if_correlated_filter']]
    assignment_params['cols_no_filter'] = cols_no_filter
    assignment_params['cols_w_filter'] = cols_w_filter
    
    return assignment_params
# ---------------------------
try:    
    student_name, student_id = myName(), myId()
except Exception as e:
    print ('You probably did not implement student-info functions, \nerror Message:',str(e))
    raise
assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"         
# ---------------------------
assignment_params = get_assignment_params(student_id)
# ---------------------------
print ('Assignment 1 (- 7 points for the test):')
print ('-----------------------')
print ('What do you need to implement?')
print ('1.   methods: myName, myId - Your personal information (- 0.5 points for the test)')
print ('2.   methods: load_dataset - load the dataset (- 0.5 points for the tests)')
print ("3.   methods: scale_fit_transform, scale_transform_for_test \n\t- the scaling type YOU NEED to implement is: " + assignment_params['scale_type'] + " (- 1.5 points for the tests)")
print ("4.a. methods: calc_distance \n\t- the distance method YOU NEED to implement is: " + assignment_params['dist_method'] + " (- 1.5 points for the tests)")
print ("4.b. methods: predict - KNN predict main flow" + " (- 2 points for the tests)")
print ("5.   methods: evaluate_performance \n\t- the evaluation metric YOU NEED to implement is: " + assignment_params['eval_metric'] + " (- 1 points for the tests)")
```

Output:
```
Assignment 1 (- 7 points for the test):
-----------------------
What do you need to implement?
1.   methods: myName, myId - Your personal information (- 0.5 points for the test)
2.   methods: load_dataset - load the dataset (- 0.5 points for the tests)
3.   methods: scale_fit_transform, scale_transform_for_test 
	- the scaling type YOU NEED to implement is: t-distribution standardization (- 1.5 points for the tests)
4.a. methods: calc_distance 
	- the distance method YOU NEED to implement is: euclidean (- 1.5 points for the tests)
4.b. methods: predict - KNN predict main flow (- 2 points for the tests)
5.   methods: evaluate_performance 
	- the evaluation metric YOU NEED to implement is: error_rate (- 1 points for the tests)
```

[Go to the beginning of the assignment](#home)
<a id="load_dataset"></a>
### Part 2 - Loading dataset and preprocess data
<img src="./images/load_dataframe.jpg" alt="load dataframe" width="100" align='left'/>

### Part 2. Instructions

<u>Part 2.</u> - load dataset from file; 
<u>method name</u>: <b>load_dataset</b><br/>

<pre>
Complete the following function to return a dataframe containing the feature vectors (X) and 
    a series containing the categories (y).
The feature vector dataframe, should not contain the column of the category. 
------------
input parameters:
- file_name - a path containing the csv file name, which you need to load the dataset from.
                  Note: this is a relative path, you should not edit this value, and use only
                        this parameter as the file name.
- category_column - the column name which represent the category column.
- index_column - the column which you should consider as the index.
   Note: set the index while loading the dataset or right after, but make sure it is before splitting 
         the dataset to X (the feature vectors) and y (the categories), which should share the same index.
------------
The return statement should look like the following:
return X, y
</pre>

<!-- grade_id: answer-load-dataset, points: 0, grade: false, solution: true -->
```python
# 2.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ANSWER 

def load_dataset(file_name, category_column, index_column):
    df_students = pd.read_csv(file_name)
    X = df_students.drop([category_column,index_column],axis=1)
    y = df_students[category_column]
    return X, y
    
```

```python
# 2.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: SELF TESTING
# Use the following code to test your implementation:
file_name = 'data' + os.sep + 'Students-Performance_numeric.csv'
student_name, student_id = myName(), myId()
assignment_params = get_assignment_params(student_id)
X, y = load_dataset(file_name, assignment_params['y_col'], assignment_params['index_col'])
# --- add additional code to check your code if needed:
# YOUR CODE HERE



```

```python
# 2.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation, if used)
# === CODE TYPE: SELF TESTING
# ---- Add your additional tests here if needed:
# YOUR CODE HERE

```

<!-- grade_id: test2-1_load_dataset, points: 0.1, grade: true, solution: false -->
```python
# 2.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 1. - Test 1 (name: test2-1_load_dataset, points: 0.1) - Sanity")
print ("\t--->Testing the implementation of 'load_dataset' ...")

file_name = 'data' + os.sep + 'Students-Performance_numeric.csv'

try:    
    student_name, student_id = myName(), myId()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)
    X, y = load_dataset(file_name, assignment_params['y_col'], assignment_params['index_col'])
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

print ("Good Job!\nYou've passed the 1st test for the 'load_dataset' function implementation :-)")
```

Output:
```
Part 1. - Test 1 (name: test2-1_load_dataset, points: 0.1) - Sanity
	--->Testing the implementation of 'load_dataset' ...
Good Job!
You've passed the 1st test for the 'load_dataset' function implementation :-)
```

<!-- grade_id: test2-2_load_dataset, points: 0.1, grade: true, solution: false -->
```python
# 2.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 1. - Test 2 (name: test2-2_load_dataset, points: 0.1) - Sanity (2)")
print ("\t--->Testing the implementation of 'load_dataset' ...")

file_name = 'data' + os.sep + 'Students-Performance_numeric.csv'

try:    
    student_name, student_id = myName(), myId()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)
    X, y = load_dataset(file_name, assignment_params['y_col'], assignment_params['index_col'])
    # ---------------------------
    X = X[assignment_params['cols_no_filter']]
    X, y = X.iloc[assignment_params['indices_1st'],:], y.iloc[assignment_params['indices_1st']]
    X, y = X.iloc[assignment_params['indices_2nd'],:], y.iloc[assignment_params['indices_2nd']]
    X_train, y_train = X.iloc[:assignment_params['num_train'],:],  y.iloc[:assignment_params['num_train']]
    X_test, y_test = X.iloc[assignment_params['num_train']:,:], y.iloc[assignment_params['num_train']:]
    # ---------------------------
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

print ("Good Job!\nYou've passed the 2nd test for the 'load_dataset' function implementation :-)")
```

Output:
```
Part 1. - Test 2 (name: test2-2_load_dataset, points: 0.1) - Sanity (2)
	--->Testing the implementation of 'load_dataset' ...
Good Job!
You've passed the 2nd test for the 'load_dataset' function implementation :-)
```

<!-- grade_id: test2-3_load_dataset, points: 0.3, grade: true, solution: false -->
```python
# 2.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 1. - Test 3 (name: test2-3_load_dataset, points: 0.3)")
print ("\t--->Testing the implementation of 'load_dataset' ...")

file_name = 'data' + os.sep + 'Students-Performance_numeric.csv'

try:    
    student_name, student_id = myName(), myId()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)
    X, y = load_dataset(file_name, assignment_params['y_col'], assignment_params['index_col'])
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

assert type(X) is pd.DataFrame or type(y) is pd.Series, "Wrong type for feature vectors dataframe or for category series"
np.testing.assert_array_equal(X.index, y.index, 'X and y should have the same index')
assert X.shape == (12000,7) and X.shape[0] == y.shape[0], 'Wrong shape forfeature vector dataframe or category series'
    
print ("Good Job!\nYou've passed the 3rd test for the 'load_dataset' function implementation :-)")
```

Output:
```
Part 1. - Test 3 (name: test2-3_load_dataset, points: 0.3)
	--->Testing the implementation of 'load_dataset' ...
Good Job!
You've passed the 3rd test for the 'load_dataset' function implementation :-)
```

[Go to the beginning of the assignment](#home)
<a id="scaling"></a>
## Part 3 -  scaling 
<img src="./images/architect-scale.jpg" alt="Scale" style="width: 500px;" align='left'/>

<b>Scaling (scale_fit_transform, scale_transform_for_test) - Permitted method usage</b>:<br />
<div class="alert alert-success">
<ul><li>&#x2714; Any function, which calculates the minimum and the maximum (such as numpy's function)</li>
<li>&#x2714; Any function, which calculates the average (such as series's function)</li>
<li>&#x2714; Any function, which calculates a standard deviation/ std/ stddev (such as dataframe's function)
    <ul><li><u>Make sure you use the sample's std</u></li></ul>
<li>&#x2714; Any function, which calculates a variance / var (such as statistic's function)</li>
    <ul><li><u>Make sure you use the sample's variance</u></li></ul>
<li>&#x2714; Any other basic mathematical function (+,*, ...) </li></ul>
</div>
<hr>
<b>Scaling (scale_fit_transform, scale_transform_for_test) - Illegal method usage</b>:<br/>
<div class="alert alert-danger">
    <ul><li>&#x2718; Any function, which performs the scaling</li></ul>
</div>
<br>
<pre>
<b><u>Important Notes - for those who need to implement (t-distribution) standardization</u></b>
  * <u>By default, numpy's methods</u> for variance (var) and standard-deviation (std) 
    <u>are on the population</u>.
  * <u>By default, pandas' methods</u> for variance (var) and standard-deviation (std)
    <u>are on the sample</u>.
</pre>         

### Part 3.a. Instructions
<u>Part 3.a. - Scaling (scale dataframe columns)</u>;
<u>method name</u>: <b>scale_fit_transform</b><br/>

<pre>
Complete the following function to return a <b>copy</b> of the train feature vector dataframe, 
       with scaled features and a dataframe containing the scaled trained information.
                   
------------
input parameter:
- X_train - the input dataframe
------------
return values:
- trained_scaling_info - a dataframe containing the trained info for each column (see the example ahead)
- X_train_scaled - a <b>copy</b> of the input X_train dataframe, with scaled columns (the scaling 
                   is dependent on the above 'scale_type', which was assigned for you above)
------------
returned statement should look as following:
return trained_scaling_info,  X_train_scaled
---------------------
</pre>
For example, let's assume we've got the following x_train:

| | sepal_length | sepal_width | petal_length | 
| :- | :- | :- | :- | 
| <b>0</b> |  5.1 | 3.5 | 1.4 |  
| <b>1</b> |  4.9 | 3 | 1.4 |   
| <b>2</b> |  4.7 | 3.2 | 1.3 |  
| <b>3</b> |  4.6 | 3.1 | 1.5 | 
| <b>4</b> |  5 | 3.6 | 1.4 | 

------

<u>Case 1</u> - the <b>'scale_type'</b> (which is assigned for you ahead equals) <b>'minmax'</b> - <br />
The trained_scaling_info output dataframe should look like this:

| | min | max |
| :- | :- | :- |
| sepal_length | 4.6 | 5.1 |
| sepal_width | 3.0 | 3.6 |
| petal_length | 1.3 | 1.5 |

And the corresponding X_train_scaled output dataframe should look like this:

| | sepal_length | sepal_width | petal_length | 
| :- | :- | :- | :- |
| <b>0</b> | 1.0 | 0.833333 | 0.5 |
| <b>1</b> | 0.6 | 0.000000 | 0.5 |
| <b>2</b> | 0.2 | 0.333333 | 0.0 |
| <b>3</b> | 0.0 | 0.166667 | 1.0 |
| <b>4</b> | 0.8 | 1.000000 | 0.5 |

------

<u>Case 2</u> - the <b>scale_type</b> </b> (which is assigned for you ahead equals) <b>'standard'</b> - <br />

| | mean | std |
| :- | :- | :- |
| sepal_length | 4.86 | 0.207364 |
| sepal_width | 3.28 | 0.258844 |
| petal_length | 1.40 | 0.070711 |
		
And the corresponding X_train_scaled output dataframe should look like this:

| | sepal_length | sepal_width | petal_length | 
| :- | :- | :- | :- |
| <b>0</b> | 1.157383 | 0.849934 | 0.000000 |
| <b>1</b> | 0.192897 |-1.081734 | 0.000000 |
| <b>2</b> | -0.771589 | -0.309067 | -1.414214 |
| <b>3</b> | -1.253831 | -0.695401 | 1.414214 |
| <b>4</b> | 0.675140 | 1.236268 | 0.000000 |

### 3. Scaling - you need to <u>implement the following method</u> ...

```python
## 3.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# === CELL TYPE: IMPLEMENTATION TASKS

try:    
    student_name, student_id = myName(), myId()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"    
    assignment_params = get_assignment_params(student_id)
except Exception as e:
    print ('You probably did not implement student-info functions, \nerror Message:',str(e))
    raise
# ----------------------------
                    
print ("The scaling type you need to implement is: " + assignment_params['scale_type'])
```

Output:
```
The scaling type you need to implement is: t-distribution standardization
```

<!-- grade_id: cell-6bb70642ef15c3b8, points: 0, grade: false, solution: true -->
```python
# 3.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ASSISTANCE TO ANSWER 
# ---- Add assistance code here IF NEEDED:
# YOUR CODE HERE

```

<!-- grade_id: answer-scale_fit_transform, points: 0, grade: false, solution: true -->
```python
# 3.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ANSWER 

def scale_fit_transform(X_train):
    X_train_scaled = X_train.copy()
    
    cols = []
    means = []
    stds = []
    
    for col in X_train.columns:
        cols.append(col)

    #Calculate mean,std useing Numpy
    for i in range(len(cols)):
        means.append(np.mean(X_train[cols[i]]))
        stds.append(np.std(X_train[cols[i]]))
    
    #Create DataFrame with t-distribution standardization
    trained_scaling_info = pd.DataFrame(list(zip(means,stds)),
    index=cols,
    columns=['mean', 'std'])

    #Scale the DataFrame X_train using t-distribution standardization scaling
    for i in range(len(X_train_scaled.columns)):
        X_train_scaled[cols[i]] = X_train_scaled[cols[i]].apply(lambda x: ((x-means[i])/stds[i]))
        
    return trained_scaling_info,  X_train_scaled
```

```python
# 3.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: SELF TESTING
# Use the following code to test your implementation:
file_name = 'data' + os.sep + 'Students-Performance_numeric.csv'
student_name, student_id = myName(), myId()
assignment_params = get_assignment_params(student_id)    
X, y = load_dataset(file_name, assignment_params['y_col'], assignment_params['index_col'])
X = X[assignment_params['cols_w_filter']]
X, y = X.iloc[assignment_params['indices_1st'],:], y.iloc[assignment_params['indices_1st']]
X, y = X.iloc[assignment_params['indices_2nd'],:], y.iloc[assignment_params['indices_2nd']]
X_train, y_train = X.iloc[:assignment_params['num_train'],:],  y.iloc[:assignment_params['num_train']]
X_test, y_test = X.iloc[assignment_params['num_train']:,:], y.iloc[assignment_params['num_train']:]
trained_scaling_info,  X_train_scaled = scale_fit_transform(X_train)
# --- add additional code to check your code if needed:
# YOUR CODE HERE


```

```python
# 3.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation, if used)
# === CODE TYPE: SELF TESTING
# ---- Add your additional tests here if needed:
# YOUR CODE HERE

```

<!-- grade_id: test3a-1_scale_fit_transform, points: 0.1, grade: true, solution: false -->
```python
# 3.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Parts 3.a. - Test 1 (name: test3a-1_scale_fit_transform, points: 0.1) - Sanity")
print ("\t--->Testing the implementation of 'scale_fit_transform' ...")

file_name = 'data' + os.sep + 'Students-Performance_numeric.csv'

try:    
    student_name, student_id = myName(), myId()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)    
    print ("\t\tTesing a '" + assignment_params['scale_type'] + "' scale type ...\n")
    X, y = load_dataset(file_name, assignment_params['y_col'], assignment_params['index_col'])
    X = X[assignment_params['cols_w_filter']]
    X, y = X.iloc[assignment_params['indices_1st'],:], y.iloc[assignment_params['indices_1st']]
    X, y = X.iloc[assignment_params['indices_2nd'],:], y.iloc[assignment_params['indices_2nd']]
    X_train, y_train = X.iloc[:assignment_params['num_train'],:],  y.iloc[:assignment_params['num_train']]
    X_test, y_test = X.iloc[assignment_params['num_train']:,:], y.iloc[assignment_params['num_train']:]
    trained_scaling_info,  X_train_scaled = scale_fit_transform(X_train)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise


print ("\nGood Job!\nYou've passed the 1st test for the 'scale_fit_transform' function implementation :-)")
```

Output:
```
Parts 3.a. - Test 1 (name: test3a-1_scale_fit_transform, points: 0.1) - Sanity
	--->Testing the implementation of 'scale_fit_transform' ...
		Tesing a 't-distribution standardization' scale type ...


Good Job!
You've passed the 1st test for the 'scale_fit_transform' function implementation :-)
```

<!-- grade_id: test3a-2_scale_fit_transform, points: 0.3, grade: true, solution: false -->
```python
# 3.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Parts 3.a. - Test 2 (name: test3a-2_scale_fit_transform, points: 0.3)")
print ("\t--->Testing the implementation of 'scale_fit_transform' ...")

file_name = 'data' + os.sep + 'Students-Performance_numeric.csv'

try:    
    student_name, student_id = myName(), myId()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)    
    print ("\t\tTesing a '" + assignment_params['scale_type'] + "' scale type ...\n")
    X, y = load_dataset(file_name, assignment_params['y_col'], assignment_params['index_col'])
    X = X[assignment_params['cols_w_filter']]
    X, y = X.iloc[assignment_params['indices_1st'],:], y.iloc[assignment_params['indices_1st']]
    X, y = X.iloc[assignment_params['indices_2nd'],:], y.iloc[assignment_params['indices_2nd']]
    X_train, y_train = X.iloc[:assignment_params['num_train'],:],  y.iloc[:assignment_params['num_train']]
    X_test, y_test = X.iloc[assignment_params['num_train']:,:], y.iloc[assignment_params['num_train']:]
    trained_scaling_info,  X_train_scaled = scale_fit_transform(X_train)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
# --------------------------- 
assert X_train_scaled.shape == X_train.shape, 'Wrong shape for scaled train dataframe'
np.testing.assert_array_equal(X_train_scaled.index, X_train.index, 'scaled train and train dataframes should have the same index')
np.testing.assert_array_equal(X_train_scaled.columns, assignment_params['cols_w_filter'], 'scaled trained and train dataframes should have the same column names')
 
print ("\nGood Job!\nYou've passed the 2nd test for the 'scale_fit_transform' function implementation :-)")
```

Output:
```
Parts 3.a. - Test 2 (name: test3a-2_scale_fit_transform, points: 0.3)
	--->Testing the implementation of 'scale_fit_transform' ...
		Tesing a 't-distribution standardization' scale type ...


Good Job!
You've passed the 2nd test for the 'scale_fit_transform' function implementation :-)
```

<!-- grade_id: test3a-3_scale_fit_transform, points: 0.3, grade: true, solution: false -->
```python
# 3.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Parts 3.a. - Test 3 (name: test3a-3_scale_fit_transform, points: 0.3)")
print ("\t--->Testing the implementation of 'scale_fit_transform' ...")

file_name = 'data' + os.sep + 'Students-Performance_numeric.csv'

try:    
    student_name, student_id = myName(), myId()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)    
    print ("\t\tTesing a '" + assignment_params['scale_type'] + "' scale type ...\n")
    X, y = load_dataset(file_name, assignment_params['y_col'], assignment_params['index_col'])
    X = X[assignment_params['cols_w_filter']]
    X, y = X.iloc[assignment_params['indices_1st'],:], y.iloc[assignment_params['indices_1st']]
    X, y = X.iloc[assignment_params['indices_2nd'],:], y.iloc[assignment_params['indices_2nd']]
    X_train, y_train = X.iloc[:assignment_params['num_train'],:],  y.iloc[:assignment_params['num_train']]
    X_test, y_test = X.iloc[assignment_params['num_train']:,:], y.iloc[assignment_params['num_train']:]
    trained_scaling_info,  X_train_scaled = scale_fit_transform(X_train)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
# --------------------------- 
if 't-distribution standardization' == assignment_params['scale_type']:
    arr_train_aggr = np.round(X_train_scaled.mean().values+0.00001,4)
    np.testing.assert_array_equal([0.,0.,0.,0.,0.], arr_train_aggr, 'Wrong average standardization scaled values')
else:    
    arr_train_aggr = np.round(X_train_scaled.min().values+0.00001,4)
    np.testing.assert_array_equal([0.,0.,0.,0.,0.], arr_train_aggr, 'Wrong average minmax scaled values')

print ("\nGood Job!\nYou've passed the 3rd test for the 'scale_fit_transform' functions implementation :-)")
```

Output:
```
Parts 3.a. - Test 3 (name: test3a-3_scale_fit_transform, points: 0.3)
	--->Testing the implementation of 'scale_fit_transform' ...
		Tesing a 't-distribution standardization' scale type ...


Good Job!
You've passed the 3rd test for the 'scale_fit_transform' functions implementation :-)
```

### Part 3.b. Instructions
<u>Part 3.b. - Scaling (scale dataframe columns)</u>;
<u>method name</u>: <b>scale_transform_for_test</b><br/>

<pre>
Complete the following function to return a <b>copy</b> of the test feature vector dataframe, 
       with scaled features.
    * The scaling will be done using the trained 'trained_scaling_info' (produced in the above 'scale_fit_transform').
input parameters:
- trained_scaling_info - a dataframe containing the trained info for each column (produced in the above 'scale_fit_transform')
- X_test - the input dataframe
------------
return values:
- X_test_scaled - a <b>copy</b> of the input test dataframe, with scaled columns (the scaling 
                   is dependent on the above 'scale_type', which was assigned for you above)
                   Note: the scaling should be done by the min/max or mean/std of each column from the training,
                         as given in the input 'trained_scaling_info' dataframe parameter.
------------
returned statement should look as following:
return X_test_scaled
---------------------
</pre>

<!-- grade_id: cell-0a5b0b729d40744f, points: 0, grade: false, solution: true -->
```python
# 3.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ASSISTANCE TO ANSWER 
# ---- Add assistance code here IF NEEDED:
# YOUR CODE HERE

```

<!-- grade_id: answer-scale_transform_for_test, points: 0, grade: false, solution: true -->
```python
# 3.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ANSWER 

def scale_transform_for_test(trained_scaling_info,  X_test):
    # YOUR CODE HERE
    X_test_scaled = X_test.copy()
    for i in X_test_scaled.columns:
        X_test_scaled[i] = (X_test_scaled[i] - trained_scaling_info.loc[i,'mean'])/trained_scaling_info.loc[i,'std']
    return X_test_scaled
```

```python
# 3.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: SELF TESTING
# Use the following code to test your implementation:
file_name = 'data' + os.sep + 'Students-Performance_numeric.csv'
student_name, student_id = myName(), myId()
assignment_params = get_assignment_params(student_id)    
X, y = load_dataset(file_name, assignment_params['y_col'], assignment_params['index_col'])
X = X[assignment_params['cols_w_filter']]
X, y = X.iloc[assignment_params['indices_1st'],:], y.iloc[assignment_params['indices_1st']]
X, y = X.iloc[assignment_params['indices_2nd'],:], y.iloc[assignment_params['indices_2nd']]
X_train, y_train = X.iloc[:assignment_params['num_train'],:],  y.iloc[:assignment_params['num_train']]
X_test, y_test = X.iloc[assignment_params['num_train']:,:], y.iloc[assignment_params['num_train']:]
trained_scaling_info,  X_train_scaled = scale_fit_transform(X_train)
X_test_scaled = scale_transform_for_test(trained_scaling_info,  X_test)
# --- add additional code to check your code if needed:
# YOUR CODE HERE

```

```python
# 3.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation, if used)
# === CODE TYPE: SELF TESTING
# ---- Add your additional tests here if needed:
# YOUR CODE HERE

```

<!-- grade_id: test3b-1_scale_transform_for_test, points: 0.1, grade: true, solution: false -->
```python
# 3.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Parts 3.b. - Test 1 (name: test3b-1_scale_transform_for_test, points: 0.1) - Sanity")
print ("\t--->Testing the implementation of 'scale_transform_for_test' ...")

file_name = 'data' + os.sep + 'Students-Performance_numeric.csv'

try:    
    student_name, student_id = myName(), myId()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)    
    print ("\t\tTesing a '" + assignment_params['scale_type'] + "' scale type ...\n")
    X, y = load_dataset(file_name, assignment_params['y_col'], assignment_params['index_col'])
    X = X[assignment_params['cols_w_filter']]
    X, y = X.iloc[assignment_params['indices_1st'],:], y.iloc[assignment_params['indices_1st']]
    X, y = X.iloc[assignment_params['indices_2nd'],:], y.iloc[assignment_params['indices_2nd']]
    X_train, y_train = X.iloc[:assignment_params['num_train'],:],  y.iloc[:assignment_params['num_train']]
    X_test, y_test = X.iloc[assignment_params['num_train']:,:], y.iloc[assignment_params['num_train']:]
    trained_scaling_info,  X_train_scaled = scale_fit_transform(X_train)
    X_test_scaled = scale_transform_for_test(trained_scaling_info,  X_test)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

print ("\nGood Job!\nYou've passed the 1st test for the 'scale_transform_for_test' function implementation :-)")
```

Output:
```
Parts 3.b. - Test 1 (name: test3b-1_scale_transform_for_test, points: 0.1) - Sanity
	--->Testing the implementation of 'scale_transform_for_test' ...
		Tesing a 't-distribution standardization' scale type ...


Good Job!
You've passed the 1st test for the 'scale_transform_for_test' function implementation :-)
```

<!-- grade_id: test3b-2_scale_transform_for_test, points: 0.3, grade: true, solution: false -->
```python
# 3.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Parts 3.b. - Test 2 (name: test3b-2_scale_transform_for_test, points: 0.3)")
print ("\t--->Testing the implementation of 'scale_transform_for_test' ...")

file_name = 'data' + os.sep + 'Students-Performance_numeric.csv'

try:    
    student_name, student_id = myName(), myId()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)    
    print ("\t\tTesing a '" + assignment_params['scale_type'] + "' scale type ...\n")
    X, y = load_dataset(file_name, assignment_params['y_col'], assignment_params['index_col'])
    X = X[assignment_params['cols_w_filter']]
    X, y = X.iloc[assignment_params['indices_1st'],:], y.iloc[assignment_params['indices_1st']]
    X, y = X.iloc[assignment_params['indices_2nd'],:], y.iloc[assignment_params['indices_2nd']]
    X_train, y_train = X.iloc[:assignment_params['num_train'],:],  y.iloc[:assignment_params['num_train']]
    X_test, y_test = X.iloc[assignment_params['num_train']:,:], y.iloc[assignment_params['num_train']:]
    trained_scaling_info,  X_train_scaled = scale_fit_transform(X_train)
    X_test_scaled = scale_transform_for_test(trained_scaling_info,  X_test)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
# --------------------------- 
assert X_test_scaled.shape == X_test.shape, 'Wrong shape for scaled test dataframe'
np.testing.assert_array_equal(X_test_scaled.index, X_test.index, 'scaled test and test dataframes should have the same index')
np.testing.assert_array_equal(X_train_scaled.columns, X_test_scaled.columns, 'scaled train and scaled test dataframes should have the same column names')

print ("\nGood Job!\nYou've passed the 2nd test for the 'scale_transform_for_test' function implementation :-)")
```

Output:
```
Parts 3.b. - Test 2 (name: test3b-2_scale_transform_for_test, points: 0.3)
	--->Testing the implementation of 'scale_transform_for_test' ...
		Tesing a 't-distribution standardization' scale type ...


Good Job!
You've passed the 2nd test for the 'scale_transform_for_test' function implementation :-)
```

<!-- grade_id: test3-scaling-train-test, points: 0.4, grade: true, solution: false -->
```python
# 3.a., 3.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Parts 3.a., 3.b. - Test 1 (name: test3-scaling-train-test, points: 0.4)")
print ("\t--->Testing the implementation of 'scale_fit_transform' and 'scale_transform_for_test' ...")

file_name = 'data' + os.sep + 'Students-Performance_numeric.csv'
try:    
    student_name, student_id = myName(), myId()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)    
    print ("\t\tTesing a '" + assignment_params['scale_type'] + "' scale type ...\n")
    X, y = load_dataset(file_name, assignment_params['y_col'], assignment_params['index_col'])
    X = X[assignment_params['cols_w_filter']]
    X, y = X.iloc[assignment_params['indices_1st'],:], y.iloc[assignment_params['indices_1st']]
    X, y = X.iloc[assignment_params['indices_2nd'],:], y.iloc[assignment_params['indices_2nd']]
    X_train, y_train = X.iloc[:assignment_params['num_train'],:],  y.iloc[:assignment_params['num_train']]
    X_test, y_test = X.iloc[assignment_params['num_train']:,:], y.iloc[assignment_params['num_train']:]
    trained_scaling_info,  X_train_scaled = scale_fit_transform(X_train)
    X_test_scaled = scale_transform_for_test(trained_scaling_info,  X_test)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
# --------------------------- 
mean_mean_scl_test =np.round(X_test_scaled.mean().values.mean(),4)
abs_mean_mean_scl_test = abs(mean_mean_scl_test)
max_mean_scl_test =np.round(X_test_scaled.mean().values.max(),4)
trained_scaling_info_means = trained_scaling_info.mean().values 

if 't-distribution standardization' == assignment_params['scale_type']:
    assert 0<abs_mean_mean_scl_test and abs_mean_mean_scl_test<0.4, 'Wrong average tandardized scaled values'
    # print (i, assignment_params['scale_type'], trained_scaling_info_means)
    np.testing.assert_array_less([14.23,3.62],np.round(trained_scaling_info_means,2), 'Wrong standardization trained_scaling_info')
else:    
    assert max_mean_scl_test>0.6, 'Wrong average minmax scaled values'
    np.testing.assert_array_less([2.1,22.3],np.round(trained_scaling_info_means,2), 'wrong minmax trained_scaling_info')

print ("\nGood Job!\nYou've passed the test for both the 'scale_fit_transform' and 'scale_transform_for_test' functions implementation :-)")
```

Output:
```
Parts 3.a., 3.b. - Test 1 (name: test3-scaling-train-test, points: 0.4)
	--->Testing the implementation of 'scale_fit_transform' and 'scale_transform_for_test' ...
		Tesing a 't-distribution standardization' scale type ...


Good Job!
You've passed the test for both the 'scale_fit_transform' and 'scale_transform_for_test' functions implementation :-)
```

## Part 4 - KNN
<img src="./images/3nn_majority.png" alt="KNN" width="300" align='left'/>

[Go to the beginning of the assignment](#home)
<a id="dist"></a>
### Part 4.a. - Calculate (vector) distance 

<b>Calculate test-train vector distance (calc_distance) - Permitted method usage</b>:<br />
<div class="alert alert-success">
<ul><li>&#x2714; Any function, which calculates the minimum and the maximum (such as numpy's function)</li>
<li>&#x2714; Any function, which calculates the average (such as series's function)</li>
<li>&#x2714; Any function, which calculates a power (such as standard python's function)</li>
<li>&#x2714; Any function, which calculates a square root / sqrt (such as statistic's function)</li>
<li>&#x2714; Any other basic vector or scalar mathematical function </li></ul>
</div>
<hr>
<b>Calculate test-train vector distance (calc_distance) - Illegal method usage</b>:<br/>
<div class="alert alert-danger">
    <ul><li>&#x2718; Any function, which performs the any of the vector distance functions</li></ul>

</div>

### Part 4.a. Instructions
<u>Part 4.a. - Vector Distance</u>;
<u>method name</u>: <b>calc_distance</b><br/>

<pre>
The following is expected:
- Change the following function to return a dataframe, containing the vector distances, 
        between each of the test vectors and each of the train vectors.
  * Each row in the returned dataframe should correspond to each of the feature vectors in the test set.
  * The index of the returned dataframe, should be equal to the indexes of 
      the test set (the X_test parameter).
  * The column names of the returned dataframe, should be equal to the indexes of 
       the train set (the X_train parameter).
  * Each cell should contain the distance between the test vector, represented by the index, 
           and train vector, represented by the column name.
  * The distance function which you need to implement, is assigned for you ahead
    ( it could be one of the following: 'manhattan', 'euclidean', 'chebyshev'
------------
input parameter:
- X_test  - test set feature vector dataframe
- X_train - train set feature vector dataframe
------------
return value:
- dist_dataframe - a dataframe containing the distances as described above.
------------
returned statement should look as following:
return dist_dataframe
---------------------
</pre>

### 4.a. Vector distance - you need to <u>implement the following method</u> ...

```python
## 4.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# === CELL TYPE: IMPLEMENTATION TASKS

try:    
    student_name, student_id = myName(), myId()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)
except Exception as e:
    print ('You probably did not implement student-info functions, \nerror Message:',str(e))
    raise
                    
print ("The distance method you need to implement is: " + assignment_params['dist_method'])
```

Output:
```
The distance method you need to implement is: euclidean
```

<!-- grade_id: cell-a3589b3c99836096, points: 0, grade: false, solution: true -->
```python
# 4.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ASSISTANCE TO ANSWER 
# ---- Add assistance code here IF NEEDED:
# YOUR CODE HERE

```

<!-- grade_id: answer-calc_distance, points: 0, grade: false, solution: true -->
```python
# 4.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ANSWER 

def calc_distance(X_test,X_train):
    testArray = []
    distance = 0
    
    #Create DataFrame for euclidean distance 
    dist_dataframe = pd.DataFrame([],
        columns = X_train.index,
        index = X_test.index)
    
    #Calculate euclidean distance for each cell on the DataFrame X_test with each cell on the DataFrame X_train
    for i in range(len(X_test)):
        calcDistRowTest = X_test.loc[X_test.index[i]]
        for j in range(len(X_train)):
            calcDistRowTrain = X_train.loc[X_train.index[j]]
            distance = euclidean_dist(calcDistRowTrain,calcDistRowTest) 
            testArray.append(distance)
        dist_dataframe.iloc[i] = testArray
        testArray.clear()
    
    return dist_dataframe
      
def euclidean_dist(x,y):
    #Calculate euclidean distance function from vectors sqrt(((y0-x0)^2) + ((y1-x1)^2) + ....)
    dist=0
    for i in range(len(x)):
        dist += (y[i] - x[i])**2
    dist = np.sqrt(dist)
    return dist  
```

```python
# 4.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: SELF TESTING
# Use the following code to test your implementation:
file_name = 'data' + os.sep + 'Students-Performance_numeric.csv'
student_name, student_id = myName(), myId()
assignment_params = get_assignment_params(student_id)
X, y = load_dataset(file_name, assignment_params['y_col'], assignment_params['index_col'])
X = X[assignment_params['cols_w_filter']]
X, y = X.iloc[assignment_params['indices_1st'],:], y.iloc[assignment_params['indices_1st']]
X, y = X.iloc[assignment_params['indices_2nd'],:], y.iloc[assignment_params['indices_2nd']]
X_train, y_train = X.iloc[:assignment_params['num_train'],:],  y.iloc[:assignment_params['num_train']]
X_test, y_test = X.iloc[assignment_params['num_train']:,:], y.iloc[assignment_params['num_train']:]
X_train_80, y_train_80 = X_train.iloc[:80,:], y_train.iloc[:80]
X_test_20, y_test_20 = X_test.iloc[:20,:], y_test.iloc[:20]
trained_scaling_info_80,  X_train_scaled_80 = scale_fit_transform(X_train_80)
X_test_scaled_20 = scale_transform_for_test(trained_scaling_info_80,  X_test_20)
dist_dataframe = calc_distance(X_test_20,X_train_80)
dist_df_scaled = calc_distance(X_test_scaled_20,X_train_scaled_80)
# --- add additional code to check your code if needed:
# YOUR CODE HERE

```

```python
# 4.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation, if used)
# === CODE TYPE: SELF TESTING
# ---- Add your additional tests here if needed:
# YOUR CODE HERE

```

<!-- grade_id: test4a-1_calc_distance, points: 0.1, grade: true, solution: false -->
```python
# 4.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 1. - Test 1 (name: test4a-1_calc_distance, points: 0.1) - Sanity")
print ("\t--->Testing the implementation of 'calc_distance' ...")

file_name = 'data' + os.sep + 'Students-Performance_numeric.csv'

try:    
    student_name, student_id = myName(), myId()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)
    print ("\t\tTesing a '" + assignment_params['dist_method'] + "' distance method ...\n")
    X, y = load_dataset(file_name, assignment_params['y_col'], assignment_params['index_col'])
    X = X[assignment_params['cols_w_filter']]
    X, y = X.iloc[assignment_params['indices_1st'],:], y.iloc[assignment_params['indices_1st']]
    X, y = X.iloc[assignment_params['indices_2nd'],:], y.iloc[assignment_params['indices_2nd']]
    X_train, y_train = X.iloc[:assignment_params['num_train'],:],  y.iloc[:assignment_params['num_train']]
    X_test, y_test = X.iloc[assignment_params['num_train']:,:], y.iloc[assignment_params['num_train']:]
    X_train_80, y_train_80 = X_train.iloc[:80,:], y_train.iloc[:80]
    X_test_20, y_test_20 = X_test.iloc[:20,:], y_test.iloc[:20]
    trained_scaling_info_80,  X_train_scaled_80 = scale_fit_transform(X_train_80)
    X_test_scaled_20 = scale_transform_for_test(trained_scaling_info_80,  X_test_20)
    dist_dataframe = calc_distance(X_test_20,X_train_80)
    dist_df_scaled = calc_distance(X_test_scaled_20,X_train_scaled_80)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise


print ("Good Job!\nYou've passed the 1st test for the 'calc_distance' function implementation :-)")
```

Output:
```
Part 1. - Test 1 (name: test4a-1_calc_distance, points: 0.1) - Sanity
	--->Testing the implementation of 'calc_distance' ...
		Tesing a 'euclidean' distance method ...

Good Job!
You've passed the 1st test for the 'calc_distance' function implementation :-)
```

<!-- grade_id: test4a-2_calc_distance, points: 0.6, grade: true, solution: false -->
```python
# 4.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 1. - Test 2 (name: test4a-2_calc_distance, points: 0.6)")
print ("\t--->Testing the implementation of 'calc_distance' ...")
file_name = 'data' + os.sep + 'Students-Performance_numeric.csv'

try:    
    student_name, student_id = myName(), myId()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)
    print ("\t\tTesing a '" + assignment_params['dist_method'] + "' distance method ...\n")
    X, y = load_dataset(file_name, assignment_params['y_col'], assignment_params['index_col'])
    X = X[assignment_params['cols_w_filter']]
    X, y = X.iloc[assignment_params['indices_1st'],:], y.iloc[assignment_params['indices_1st']]
    X, y = X.iloc[assignment_params['indices_2nd'],:], y.iloc[assignment_params['indices_2nd']]
    X_train, y_train = X.iloc[:assignment_params['num_train'],:],  y.iloc[:assignment_params['num_train']]
    X_test, y_test = X.iloc[assignment_params['num_train']:,:], y.iloc[assignment_params['num_train']:]
    X_train_80, y_train_80 = X_train.iloc[:80,:], y_train.iloc[:80]
    X_test_20, y_test_20 = X_test.iloc[:20,:], y_test.iloc[:20]
    trained_scaling_info_80,  X_train_scaled_80 = scale_fit_transform(X_train_80)
    X_test_scaled_20 = scale_transform_for_test(trained_scaling_info_80,  X_test_20)
    dist_dataframe = calc_distance(X_test_20,X_train_80)
    dist_df_scaled = calc_distance(X_test_scaled_20,X_train_scaled_80)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
assert dist_dataframe.shape == dist_df_scaled.shape, 'Mismatch shape of scaled and non scaled distance dataframes'
np.testing.assert_array_equal(dist_dataframe.index, X_test_scaled_20.index, 'Index of distance dataframe should match test index')
np.testing.assert_array_equal(dist_dataframe.columns, X_train_80.index, 'Column names of distance dataframe should match training index')
np.testing.assert_array_equal(dist_df_scaled.columns, X_train_80.index, 'Column names of distance dataframe should match training index')

print ("Good Job!\nYou've passed the 2nd test for the 'calc_distance' function implementation :-)")
```

Output:
```
Part 1. - Test 2 (name: test4a-2_calc_distance, points: 0.6)
	--->Testing the implementation of 'calc_distance' ...
		Tesing a 'euclidean' distance method ...

Good Job!
You've passed the 2nd test for the 'calc_distance' function implementation :-)
```

<!-- grade_id: test4a-3_calc_distance, points: 0.8, grade: true, solution: false -->
```python
# 4.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 1. - Test 3 (name: test4a-3_calc_distance, points: 0.8)")
print ("\t--->Testing the implementation of 'calc_distance' ...")
print ("\n\t====> Full grading test - the following test is invisible. You will not know if you passed it, before grades are returned to the students")

```

Output:
```
Part 1. - Test 3 (name: test4a-3_calc_distance, points: 0.8)
	--->Testing the implementation of 'calc_distance' ...

	====> Full grading test - the following test is invisible. You will not know if you passed it, before grades are returned to the students
```

[Go to the beginning of the assignment](#home)
<a id="predict"></a>
### Part 4.b. - KNN predict - main flow

<b>KNN - main flow (predict) - Permitted method usage</b>:<br />
<div class="alert alert-success">
<ul><li>&#x2714; use the 'calc_distance' function, implemented above </li>
<li>&#x2714; Any function, which calculates maximum, top scores, sorting</li>
<li>&#x2714; Any function, which calculates the combination of the above with arg... (such as argmax)</li>
<li>&#x2714; Any function, which calculates a square root / sqrt (such as statistic's function)</li>
<li>&#x2714; Any other basic vector or scalar mathematical function </li></ul>
</div>
<hr>
<b>KNN - main flow (predict) - Illegal method usage</b>:<br/>
<div class="alert alert-danger">
    <ul><li>&#x2718; Any function, which performs the main or partial flow of KNN</li></ul>

</div>

### Part 4.b. Instructions
<u>Part 4.b.</u> - KNN - predict;
<u>method name</u>: <b>predict</b><br/>

<pre>
The following is expected:
- Change the following function to return a series consisting of the predicted category series. 
   * to do so, use the 'calc_distance' function, implemented above, to get the distance, between 
     the feature vector, to each of the training feature vectors.
------------
input parameter:
- X_test - the test set feature vector dataframe
- X_train - the train set feature vector dataframe
- y_train - the train set category series 
- k - the number of Nearest Neighbors (we select the majority out of k votes)
------------
return value:
- y_predicted - the predicted category series, for each test set example, determined by the KNN classifier.
                  * Note: the returned predicted category series, should contain an index, which
                          equals to the X_test index.
------------
returned statement should look as following:
return y_predicted
---------------------
</pre>

<!-- grade_id: cell-fd933b7957439b5d, points: 0, grade: false, solution: true -->
```python
# 4.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ASSISTANCE TO ANSWER
# ---- Add assistance code here IF NEEDED:
# YOUR CODE HERE

```

<!-- grade_id: answer-predict, points: 0, grade: false, solution: true -->
```python
# 4.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ANSWER 

def predict(X_test, X_train, y_train, k):
    
    #Calculate euclidean distance
    dfDistToPredicte = calc_distance(X_test,X_train)
    
    sortedRowList = []
    listOfTrainsMacthId = []
    calcPredicte = []
    
    ''' 
    Get all the of the sorted cells
    and sort it one row by one , each row in X_test is added to getRowSorted and
    then add it to sortedRowList with the index match on X_train as sorted row
    '''
    for ind in dfDistToPredicte.index:
        dfDistToPredicte.loc[ind].sort_values
        getRowSorted=dfDistToPredicte.T.sort_values(by=ind).T
        sortedRowList.append(getRowSorted.loc[ind])

    calcPredicte=getKNeighbors(len(X_test),sortedRowList,k)
    
    #get predicted
    for ind in calcPredicte:
        listOfTrainsMacthId.append(y_train.loc[ind])
    
    #Create Series to preticed results
    y_predicted = pd.Series(0,index=X_test.index)
    
    #Get the predicted from the categories
    for ind in range(len(X_test)):
        predicted=0
        for jnd in range(k):
            predicted += listOfTrainsMacthId[ind].values[jnd] 
        if (predicted < k/2):
            y_predicted.iloc[ind] = 0
        else:
            y_predicted.iloc[ind] = 1
    
    return y_predicted

def getKNeighbors(amount,lists,k):
    kNeighbors = []
    calcPredicte = []
    #Get the K Neighbors
    for i in range(amount):
        for j in range(k):
            kNeighbors.append(lists[i].index[j])
        #add the K Neighbors to 
        calcPredicte.append(kNeighbors.copy())
        kNeighbors.clear()
    return calcPredicte
    
    

```

```python
# 4.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: SELF TESTING
# Use the following code to test your implementation:
file_name = 'data' + os.sep + 'Students-Performance_numeric.csv'
student_name, student_id = myName(), myId()
assignment_params = get_assignment_params(student_id)
X, y = load_dataset(file_name, assignment_params['y_col'], assignment_params['index_col'])
X = X[assignment_params['cols_w_filter']]
X, y = X.iloc[assignment_params['indices_1st'],:], y.iloc[assignment_params['indices_1st']]
X, y = X.iloc[assignment_params['indices_2nd'],:], y.iloc[assignment_params['indices_2nd']]
X_train, y_train = X.iloc[:assignment_params['num_train'],:],  y.iloc[:assignment_params['num_train']]
X_test, y_test = X.iloc[assignment_params['num_train']:,:], y.iloc[assignment_params['num_train']:]
X_train_100, y_train_100 = X_train.iloc[100:200,:], y_train.iloc[100:200]
X_test_20, y_test_20 = X_test.iloc[20:40,:], y_test.iloc[20:40]
trained_scaling_info_100,  X_train_scaled_100 = scale_fit_transform(X_train_100)
X_test_scaled_20 = scale_transform_for_test(trained_scaling_info_100,  X_test_20)
y_predicted = predict(X_test_20, X_train_100, y_train_100, 1)
y_predicted_scaled = predict(X_test_scaled_20, X_train_scaled_100, y_train_100, 1)
# --- add additional code to check your code if needed:
# YOUR CODE HERE

```

```python
# 4.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation, if used)
# === CODE TYPE: SELF TESTING
# ---- Add your additional tests here if needed:
# YOUR CODE HERE

```

<!-- grade_id: test4b-1_predict, points: 0.1, grade: true, solution: false -->
```python
# 4.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 1. - Test 1 (name: test4b-1_predict, points: 0.1) - Sanity")
print ("\t--->Testing the implementation of 'predict' ...")

file_name = 'data' + os.sep + 'Students-Performance_numeric.csv'

try:    
    student_name, student_id = myName(), myId()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)
    X, y = load_dataset(file_name, assignment_params['y_col'], assignment_params['index_col'])
    X = X[assignment_params['cols_w_filter']]
    X, y = X.iloc[assignment_params['indices_1st'],:], y.iloc[assignment_params['indices_1st']]
    X, y = X.iloc[assignment_params['indices_2nd'],:], y.iloc[assignment_params['indices_2nd']]
    X_train, y_train = X.iloc[:assignment_params['num_train'],:],  y.iloc[:assignment_params['num_train']]
    X_test, y_test = X.iloc[assignment_params['num_train']:,:], y.iloc[assignment_params['num_train']:]
    X_train_100, y_train_100 = X_train.iloc[100:200,:], y_train.iloc[100:200]
    X_test_20, y_test_20 = X_test.iloc[20:40,:], y_test.iloc[20:40]
    trained_scaling_info_100,  X_train_scaled_100 = scale_fit_transform(X_train_100)
    X_test_scaled_20 = scale_transform_for_test(trained_scaling_info_100,  X_test_20)
    y_predicted = predict(X_test_20, X_train_100, y_train_100, 1)
    y_predicted_scaled = predict(X_test_scaled_20, X_train_scaled_100, y_train_100, 1)
except Exception as e:
    print ('You probably did not implement student-info functions, \nerror Message:',str(e))
    raise
# ---------------------------

print ("Good Job!\nYou've passed the 1st test for the 'predict' function implementation :-)")
```

Output:
```
Part 1. - Test 1 (name: test4b-1_predict, points: 0.1) - Sanity
	--->Testing the implementation of 'predict' ...
Good Job!
You've passed the 1st test for the 'predict' function implementation :-)
```

<!-- grade_id: test4b-2_predict, points: 0.9, grade: true, solution: false -->
```python
# 4.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 4.b. - Test 2 (name: test4b-2_predict, points: 0.9)")
print ("\t--->Testing the implementation of 'predict' ...")

file_name = 'data' + os.sep + 'Students-Performance_numeric.csv'

try:    
    student_name, student_id = myName(), myId()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)
    X, y = load_dataset(file_name, assignment_params['y_col'], assignment_params['index_col'])
    X = X[assignment_params['cols_w_filter']]
    X, y = X.iloc[assignment_params['indices_1st'],:], y.iloc[assignment_params['indices_1st']]
    X, y = X.iloc[assignment_params['indices_2nd'],:], y.iloc[assignment_params['indices_2nd']]
    X_train, y_train = X.iloc[:assignment_params['num_train'],:],  y.iloc[:assignment_params['num_train']]
    X_test, y_test = X.iloc[assignment_params['num_train']:,:], y.iloc[assignment_params['num_train']:]
    X_train_100, y_train_100 = X_train.iloc[100:200,:], y_train.iloc[100:200]
    X_test_20, y_test_20 = X_test.iloc[20:40,:], y_test.iloc[20:40]
    trained_scaling_info_100,  X_train_scaled_100 = scale_fit_transform(X_train_100)
    X_test_scaled_20 = scale_transform_for_test(trained_scaling_info_100,  X_test_20)
    y_predicted = predict(X_test_20, X_train_100, y_train_100, 3)
    y_predicted_scaled = predict(X_test_scaled_20, X_train_scaled_100, y_train_100, 3)
except Exception as e:
    print ('You probably did not implement student-info functions, \nerror Message:',str(e))
    raise
# ---------------------------
assert y_predicted.shape==y_predicted_scaled.shape and y_test_20.shape==y_predicted.shape,'Wrong shape for predicted series'
np.testing.assert_array_equal(y_predicted.index, y_test_20.index, 'Index of predicted series should match test series')
np.testing.assert_array_equal(y_predicted_scaled.index, y_test_20.index, 'Index of predicted series should match test series')
np.testing.assert_equal(np.any(np.not_equal(y_predicted_scaled.values,y_test_20.values)), True,'predicted and test series should not be identicle')
print ("Good Job!\nYou've passed the 2nd test for the 'predict' function implementation :-)")
```

Output:
```
Part 4.b. - Test 2 (name: test4b-2_predict, points: 0.9)
	--->Testing the implementation of 'predict' ...
Good Job!
You've passed the 2nd test for the 'predict' function implementation :-)
```

<!-- grade_id: test4b-3_predict, points: 1, grade: true, solution: false -->
```python
# 4.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 4.b. - Test 3 (name: test4b-3_predict, points: 1)")
print ("\t--->Testing the implementation of 'predict' ...")
print ("\n\t====> Full grading test - the following test is invisible. You will not know if you passed it, before grades are returned to the students")


```

Output:
```
Part 4.b. - Test 3 (name: test4b-3_predict, points: 1)
	--->Testing the implementation of 'predict' ...

	====> Full grading test - the following test is invisible. You will not know if you passed it, before grades are returned to the students
```

[Go to the beginning of the assignment](#home)
<a id="evaluation"></a>
### Part 5 - Evaluation
<img src="./images/accuracy.jpg" alt="Accuracy" width="300" align='left'/>

<b>KNN - main flow (predict) - Permitted method usage</b>:<br />
<div class="alert alert-success">
<li>&#x2714; Any other basic vector or scalar mathematical function </li></ul>
</div>
<hr>
<b>KNN - main flow (predict) - Illegal method usage</b>:<br/>
<div class="alert alert-danger">
    <ul><li>&#x2718; Any function, which calculates a confusion matrix</li>
    <li>&#x2718; Any function, which calculates percision, recall, accuracy or error rate</li></ul>
</div>

### Part 5. Instructions
<u>Part 5 - Evaluation </u>; 
<u>method name</u>: <b>evaluate_performance</b><br/>

<pre>
The following is expected:
- the method needs to return the accuracy (correctly predicted ratio)
------------
input parameter:
- y_predicted -  a series of containing all predicted class values per test instance
- y_test - a series of containing all actual class values per test instance
------------
return value:
------------
- evaluate_value - the value calculated by the evaluation metric assigned for you ahead
<br>
<b><u>Important Note - for those who need to implement precision or recall</u></b>
  * <u>The <b>positive</b> category value</u> refers to <u>the value <b>1</b></u>.
  * <u>The <b>negative</b> category value</u> refers to <u>the value <b>0</b></u>.
</pre>         

### 5. Evaluation - you need to <u>implement the following method</u> ...

```python
## 5.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# === CELL TYPE: IMPLEMENTATION TASKS

try:    
    student_name, student_id = myName(), myId()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)
except Exception as e:
    print ('You probably did not implement student-info functions, \nerror Message:',str(e))
    raise
# ---------------------------
                    
print ("The evaluation metric you need to implement is: " + assignment_params['eval_metric'])
```

Output:
```
The evaluation metric you need to implement is: error_rate
```

<!-- grade_id: cell-6909f244b4cc4cb3, points: 0, grade: false, solution: true -->
```python
# 5.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ASSISTANCE TO ANSWER
# ---- Add assistance code here IF NEEDED:
# YOUR CODE HERE

```

<!-- grade_id: answer-evaluate_performance, points: 0, grade: false, solution: true -->
```python
# 5.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ANSWER 

def evaluate_performance(y_predicted,y_test):
    TP = FN = FP = TN = 0
    
    for i in range(len(y_test)):
        if y_predicted.iloc[i] == 1 and y_test.iloc[i] == 1:
            TP += 1
        elif y_predicted.iloc[i] == 1 and y_test.iloc[i] == 0:
            FP += 1
        elif y_predicted.iloc[i] == 0 and y_test.iloc[i] == 1:
            FN += 1
        elif y_predicted.iloc[i] == 0 and y_test.iloc[i] == 0:
            TN += 1
            
    evaluate_value = (FP+FN)/(TP+FP+FN+TN)
    
    return evaluate_value

```

```python
# 5.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: SELF TESTING
# Use the following code to test your implementation:
file_name = 'data' + os.sep + 'Students-Performance_numeric.csv'
student_name, student_id = myName(), myId()
assignment_params = get_assignment_params(student_id)
indx = [3,1,6,5,7,44,32,14,10,11]
y_pred = pd.Series([1,1,1,0,1,0,1,0,1,0],index=indx)
y_test = pd.Series([0,0,1,1,1,0,0,0,1,0],index=indx)
eval_res = evaluate_performance(y_pred, y_test)
# --- add additional code to check your code if needed:
# YOUR CODE HERE

```

```python
# 5.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation, if used)
# === CODE TYPE: SELF TESTING
# ---- Add your additional tests here if needed:
# YOUR CODE HERE

```

<!-- grade_id: test5-1_evaluate_performance, points: 0.1, grade: true, solution: false -->
```python
# 5.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 1. - Test 1 (name: test5-1_evaluate_performance, points: 0.1) - Sanity")
print ("\t--->Testing the implementation of 'evaluate_performance' ...")

file_name = 'data' + os.sep + 'Students-Performance_numeric.csv'

try:    
    student_name, student_id = myName(), myId()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)
    print ("\t\tTesing a '" + assignment_params['eval_metric'] + "' evaluation method ...\n")
    indx = [3,1,6,5,7,44,32,14,10,11]
    y_pred = pd.Series([1,1,1,0,1,0,1,0,1,0],index=indx)
    y_test = pd.Series([0,0,1,1,1,0,0,0,1,0],index=indx)
    eval_res = evaluate_performance(y_pred, y_test)
    
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

print ("Good Job!\nYou've passed the 1st test for the 'evaluate_performance' function implementation :-)")
```

Output:
```
Part 1. - Test 1 (name: test5-1_evaluate_performance, points: 0.1) - Sanity
	--->Testing the implementation of 'evaluate_performance' ...
		Tesing a 'error_rate' evaluation method ...

Good Job!
You've passed the 1st test for the 'evaluate_performance' function implementation :-)
```

<!-- grade_id: test5-2_evaluate_performance, points: 0.9, grade: true, solution: false -->
```python
# 5.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 2. - Test 2 (name: test5-2_evaluate_performance, points: 0.9)")
print ("\t--->Testing the implementation of 'evaluate_performance' ...")

file_name = 'data' + os.sep + 'Students-Performance_numeric.csv'

try:    
    student_name, student_id = myName(), myId()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)
    print ("\t\tTesing a '" + assignment_params['eval_metric'] + "' evaluation method ...\n")
    indx = [3,1,6,5,7,44,32,14,10,11]
    y_pred = pd.Series([1,1,1,0,1,0,1,0,1,0],index=indx)
    y_test = pd.Series([0,0,1,1,1,0,0,0,1,0],index=indx)
    eval_res = evaluate_performance(y_pred, y_test)
    
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
    
if assignment_params['eval_metric'] == 'precision':
    assert eval_res == 0.5, 'Wrong precision value'
elif assignment_params['eval_metric'] == 'recall':
    assert eval_res == 0.75, 'Wrong recall value'
elif assignment_params['eval_metric'] == 'accuracy':
    assert eval_res == 0.6, 'Wrong accuracy value'
else: # if assignment_params['eval_metric'] == 'error_rate':
    assert eval_res == 0.4, 'Wrong error-rate value'
    
print ("Good Job!\nYou've passed the 2nd test for the 'evaluate_performance' function implementation :-)")
```

Output:
```
Part 2. - Test 2 (name: test5-2_evaluate_performance, points: 0.9)
	--->Testing the implementation of 'evaluate_performance' ...
		Tesing a 'error_rate' evaluation method ...

Good Job!
You've passed the 2nd test for the 'evaluate_performance' function implementation :-)
```

```python

```

```python

```

```python

```