<a id="home"></a>
## Assignment 2
In this assignment you will need to implement functions throughout the regression flow.<br/>

To receive YOUR specific task, go to section 1c - [Your personal implementation tasks](#1c), 
<b>after</b> you run the preceding step - [import packages](#Preceding) and implement & run the [myName](#myName), [myId](#myId) functions. <br/>

<b>Index of functions you need to implement (total possible points: 9</b>:<br/>

| Section | Section-name | points | Section | Section-name | points | Section | Section-name | points |
| :- | :- | :- | :- | :- | :- | :- | :- | :- | 
| Guidelines | [General assignment instructions](#gen_instructions) | | Imports | [Import packages](#Preceding) | | Part 1. | [Student info](#student_info) | 0.5 |
| 1c. | [info: YOUR personal implementation tasks](#1c) | | The dataset | [Dataset description](#about_dataset) |  | Part 2. | [Load the dataset (for linear regression)](#load_dataset) | 1 |
| Part 3. | [Scaling](#scaling) | 1.5 | Part 4. | [Linear Regression - predict](#predict) | 1 | Part 5. | [Training - fit](#fit) | 3 |
| Part 6. | [Model Evaluation](#evaluation) | 2 | 

<a id="gen_instructions"></a>
## Working on the assignment, todo and not todo:
<b><u>To-do</u></b>:<br />
<div class="alert alert-success">
<ul><li>&#x2714; <b> Work by yourself</b> and submit your own assignment, <b>no pairing</b> to other students</li>
<li>&#x2714; Test and save your assignment - <b>submit the last tested and saved version</b></li>
<li>&#x2714; Work on the <b>original notebook</b> of the assignment</li>
<li>&#x2714; submit <b>only</b> the <b>notebook (.ipynb)</b> file</li>
    <ul><li>&#x2714;<b>Keep original name</b>: <u>Assignment2.ipynb</u></li></ul>        
</div>
<b><u>Not to-do</u></b>:<br />
<div class="alert alert-danger">
<ul><li>&#x2718; do <b>NOT</b> submit an <b>empty assignment</b></li>
<li>&#x2718; do <b>NOT submit extra files</b>, <u>unless you're asked to do so</u></li>
<li>&#x2718; <b>Do NOT submit a .py/.txt/.rar/.zip (or any non (.ipynb) file)</b> <u>version for the notebook</u> of the assignment</li>
<li>&#x2718; <b>do NOT change</b> the notebook file name</li></ul>
</div>

<b>Not good - Any of the following</b>:<br/>
<div class="alert alert-danger">
<ul><li>&#x2718; Code not Implemented</li>
<li>&#x2718; No returned value</li>
<li>&#x2718; Did not remove the '<code>raise NotImplementedError()</code>' line:</li></ul><br/>
<code>
# ------------ >>>>>>>> RUN THIS CODE CELL <<<<<<<< ------------
# --------  (run after implementation) 
# === CODE TYPE: ANSWER
def method_to_implement():
    # YOUR CODE HERE
    raise NotImplementedError()
</code></div>

<b>Good - All of the following</b>:<br/>
<div class="alert alert-success">
<ul>
<li>&#x2714; Your code starts one line after the '<code># YOUR CODE HERE</code>' remark</li>
<li>&#x2714; Removed the '<code>raise NotImplementedError()</code>' line</li>
<li>&#x2714; Implemented code</li>
<li>&#x2714; returned expected value</li></ul><br/>
<code>
# ------------ >>>>>>>> RUN THIS CODE CELL <<<<<<<< ------------
# --------  (run after implementation) 
# === CODE TYPE: ANSWER
def method_to_implement():
    # YOUR CODE HERE
    implemented code ...
</code></div>

[Go to the beginning of the assignment](#home)
## Preceding Step - import modules (packages)
This step is necessary in order to use external modules (packages). <br/>
**No other modules are allowed**

**Use the following libraries for the assignment, when needed**:

<b>Import packaged - permitted packages</b>:<br />
<div class="alert alert-success">
&#x2714; Any package or module imported in the following code cell are permitted
</div>
<hr>
<b>Import packaged - Illegal packages</b>:<br/>
<div class="alert alert-danger">
&#x2718; Any package or module NOT IMPORTED in the following code cell are ILLEGAL
</div>

```python
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# === CELL TYPE: IMPORTS AND SETUP 

import sys                       # for testing use only
import os                        # for testing use only
from datetime import datetime    # for testing use only
import hashlib                   # for testing use only
import random                    # for testing use only
import pathlib 
import pandas as pd
import numpy as np
import math
import statistics
import matplotlib.pyplot as plt
from mpl_toolkits.mplot3d import axes3d, Axes3D
get_ipython().run_line_magic('matplotlib', 'inline')
```

[Go to the beginning of the assignment](#home)
<a id="student_info"></a>
<h3>Part 1 - Student information methods and YOUR personal implementation:</h3><br/>
The following 2 cells consist of 2 methods which aim to validate your details.<br/>
Please <b>make sure to implement <u>both</u> of them</b><br/><br/>
<table align='left'>
    <tr><td>
        <img src="./images/id-card.png" alt="name and id" width="150" align='left'/>
    </td></tr>
</table>

<b>Part 1.a. - Student information - <u>make sure</u></b>:<br />
<div class="alert alert-success">
&#x2714; Make sure: <u>ID number</u> returned here (the output of myId()) <i>MATCHES</i> <u>ID number in Moodle</u><br />
&#x2714; Make sure: <u>ID number</u> returned here (the output of myId()) is an <u>int number</u><br />
&#x2714; Make sure: <u>Student name</u> returned here (the output of myName()) is in <u>English</u> and is a <u>string</u><br />
</div>
<hr>
<b>Student information - <u>avoid</u> (make sure you <u>DO NOT do the following</u>)</b>:<br/>
<div class="alert alert-danger">
&#x2718; Avoid: <u>ID number</u> returned here (the output of myId()) <i>DIFFERENT</i> from <u>ID number in Moodle</u><br/>
&#x2718; Avoid: <u>ID number</u> returned here (the output of myId()) as a <u>string</u> or a <u>float</u><br/>
</div>

### Part 1.a. Instructions
<a id="myName"></a>
<u>Part 1.a.</u> Student information methods; 
<u>method name</u>: <b>myName</b><br/>
<pre>
Complete the following function to return your full name (as a string)
    For example: assume your name is John Smith, you should write:
    return 'John Smith'
------------
return value:
- your full name (as a string)
---------------------
</pre>

<!-- grade_id: answer-my_name, points: 0, grade: false, solution: true -->
```python
# 1.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ANSWER 

def myName():
    return 'Roi Aloosh'
```

```python
# 1.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: SELF TESTING
# Use the following code to test your implementation:
# ~~~~ CALL YOUR 'myName' IMPLEMENTATION: ~~~~ #
student_name = myName()
# --- add additional code to check your code if needed:
# YOUR CODE HERE

```

<a id="myId"></a>
<u>Part 1.b.</u> Student information methods</u>; 
<u>method name</u>: <b>myId</b><br/>
<pre>
Complete the following function to return your ID number (as an integer number)
    For example: assume your ID number is 1234, you should write:
    return 1234
------------
return value:
- idnumber your ID number (as an integer number)
------------
The return statement should look like the following:
return idnumber
</pre>

<!-- grade_id: answer-my_id, points: 0, grade: false, solution: true -->
```python
# 1.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ANSWER 

def myId():
    return 313233686
```

```python
# 1.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: SELF TESTING
# Use the following code to test your implementation:
# ~~~~ CALL YOUR 'myId' IMPLEMENTATION: ~~~~ #
student_id = myId()
# --- add additional code to check your code if needed:
# YOUR CODE HERE

```

<!-- grade_id: test1-1_student_info, points: 0.5, grade: true, solution: false -->
```python
# 1.a. 1.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 1.a. 1.b. - Test 1 (name: test1-1_student_info, points: 0.5)")
print ("\t--->Testing the implementation of 'myName' and 'myId' ...")

# dataframe for output:
dt1 = datetime.now()
try:
    student_name, student_id = myName(), myId()
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"

s_datetime =  [dt1.strftime('%Y-%m-%d %H:%M:%S'), datetime.now().strftime('%Y-%m-%d %H:%M:%S')]
df_output = pd.DataFrame({'value': [student_name, student_id], 'type': [type(student_name), type(student_id)], 'date_time': s_datetime},index=['student_name','student_id'])

print ("\nGood Job!\nYou've passed the test for the implementation of 'myName' and 'myId'  :-)")

print ('\nOutput dataframe:')
print ('-----------------')
df_output

```

Output:
```
Part 1.a. 1.b. - Test 1 (name: test1-1_student_info, points: 0.5)
	--->Testing the implementation of 'myName' and 'myId' ...

Good Job!
You've passed the test for the implementation of 'myName' and 'myId'  :-)

Output dataframe:
-----------------
                   value           type            date_time
student_name  Roi Aloosh  <class 'str'>  2021-04-22 12:57:32
student_id     313233686  <class 'int'>  2021-04-22 12:57:32
```

[Go to the beginning of the assignment](#home)
<a id="1c"></a>
## 1.c. Info: Your personal implementation tasks - <u>implement the following methods</u> ...
Run the following code cell to and you will know what do YOU NEED to implement

```python
# 1.c.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# === CELL TYPE: IMPLEMENTATION TASKS

# --------------------------------------------------------------- #
# Run the following and get information about your specific tasks #
# --------------------------------------------------------------- #

def get_assignment_params(student_id):
    num_modulo = 2 ** 31-1
    created_1st_state = int(hashlib.md5(str(student_id).encode('utf-8')).hexdigest(),16) % num_modulo
    created_2nd_state = int(hashlib.md5(str(created_1st_state).encode('utf-8')).hexdigest(),16) % num_modulo
    created_3rd_state = int(hashlib.md5(str(created_2nd_state).encode('utf-8')).hexdigest(),16) % num_modulo
    
    scale_types, eval_metrics = ['minmax normalization', 't-distribution standardization'], ['SAE', 'MAE','RMSE']
    
    assignment_params = {}
    
    assignment_params['cols'] = ['x_0', 'year', 'is_diesel', 'owner', 'is_automatic', 'km_driven', 'seats', 'mileage (kmpl)', 'max_horse_power', 'is_individual_seller', 'engine_capacity_(cc)']
    assignment_params['y_col'], assignment_params['fictive_attr'] = 'selling_price', 'x_0'
    
    random.seed(created_3rd_state)  
    assignment_params['scale_type'] = random.choice(scale_types)
    
    random.seed(created_2nd_state)  
    assignment_params['eval_metric'] = random.choice(eval_metrics)
    
    return assignment_params
# ---------------------------
try:    
    student_name, student_id = myName(), myId()
except Exception as e:
    print ('You probably did not implement student-info functions, \nerror Message:',str(e))
    raise
assert type(student_name) is str and type(student_id) is int, "name is not a string or id is not an integer"         
# ---------------------------
assignment_params = get_assignment_params(student_id)
# ---------------------------
print ('Assignment 2 (- 9 points for the test):')
print ('-----------------------')
print ('What do you need to implement?')
print ('1. methods: myName, myId - Your personal information (- 0.5 points for the test)')
print ('2. method: load_dataset - load the dataset (- 1 point for the tests)')
print ("3. methods: scale_fit_transform, scale_transform_for_test \n\t- the scaling type YOU NEED to implement is: " + assignment_params['scale_type'] + " (- 1.5 points for the tests)")
print ("4. Linear Regression - predict (- 1 points for the test)")
print ("5. Training - fit (- 3 points for the test)")
print ("6.method: evaluate_regression, calc_r_squared (- 2 points for the tests)\n\t- Note for 'evaluate_regression' - the evaluation metric YOU NEED to implement is: " + assignment_params['eval_metric'] + "")
```

Output:
```
Assignment 2 (- 9 points for the test):
-----------------------
What do you need to implement?
1. methods: myName, myId - Your personal information (- 0.5 points for the test)
2. method: load_dataset - load the dataset (- 1 point for the tests)
3. methods: scale_fit_transform, scale_transform_for_test 
	- the scaling type YOU NEED to implement is: t-distribution standardization (- 1.5 points for the tests)
4. Linear Regression - predict (- 1 points for the test)
5. Training - fit (- 3 points for the test)
6.method: evaluate_regression, calc_r_squared (- 2 points for the tests)
	- Note for 'evaluate_regression' - the evaluation metric YOU NEED to implement is: RMSE
```

[Go to the beginning of the assignment](#home)
<a id="about_dataset"></a>
### About the dataset:
The following is an explanation of the columns in the dataset, including the target:
* year - manufactured year
* selling_price	- target value, price of the car
* is_diesel - binary feature, is the fuel type 'diesel'?
* owner - what is the ownership (1 for 1st, 2 for 2nd and so on)
* is_automatic - binary feature, is it an 'automatic' gear?
* km_driven - how many k.m. were driven before the car was sold.
* seats - number of seats in the car
* mileage (kmpl) - fuel consumption.
* max_horse_power 
* is_individual_seller - binary feature, is the seller an 'individual' seller?	
* engine_capacity_(cc)

[Go to the beginning of the assignment](#home)
<a id="load_dataset"></a>
### Part 2 - Loading dataset  (for linear regression)
<img src="./images/load_dataframe.jpg" alt="load dataframe" width="100" align='left'/>

### Part 2. Instructions
<u>Part 2.</u> - load dataset from file; 
<u>method name</u>: <b>load_dataset</b><br/>

<pre>
Complete the following function to return a dataframe containing the feature vectors (X) and 
    a series containing the targets (y).
    
The feature vector dataframe, should not contain the column of the target. 

It is expected the the X will contain a fictive column of features, with 
 the constant value of 1, representing the x_0 fictive feature
   * The column name should be called 'x_0'.
   * The 'x_0' column name should be placed as the first column (before any other columns).
------------
input parameters:
- file_name - a path containing the csv file name, which you need to load the dataset from.
                  Note: this is a relative path, you should not edit this value, and use only
                        this parameter as the file name.
- target_column - the column name which represent the target column.
- fictive_attr_name - the name of the fictive feature column, with the constant value of 1, 
                      representing the x_0 fictive feature.
------------
The return statement should look like the following:
return X, y
</pre>
For example, `let's assume` we've received the following input parameters:
* `file_name=...` - some file name
* <code>target_column='price_($k)'</code>
* <code>fictive_attr_name='x_0'</code>

------

<b><u>Step 1</u></b>: Let's assume the following <b>dataframe was loaded from the csv</b> file:<br />

| | size | #bedrooms | price_($k) |
| :- | :- | :- | :- |
| 0 | 2104 | 5 | 460 |
| 1 | 1416 | 3 | 232 |
| 2 | 852 | 2 | 178 |
 
------

<br /><b><u>Step 2</u>: After adding the fictive parameter it should look like this</b>:    

| | x_0 | size | #bedrooms | price_($k) |
| :- | :- | :- | :- | :- |
| 0 | 1 | 2104 | 5 | 460 |
| 1 | 1 | 1416 | 3 | 232 |
| 2 | 1 | 852 | 2 | 178 |    

------

<br /><b><u>Step 3</u>: We should return a feature vector dataframe (X) looking like this</b>:

| | x_0 | size | #bedrooms |
| :- | :- | :- | :- |
| 0 | 1 | 2104 | 5 |
| 1 | 1 | 1416 | 3 |
| 2 | 1 | 852 | 2 |

<b>And a target series (y) looking like this</b>:
    
| | |
| :- | :- |
| 0 | 460 |
| 1 | 232 |
| 2 | 178 |

-----

<b>The return statement should look like this</b>:<br />
return X, y

<!-- grade_id: answer-load_dataset, points: 0, grade: false, solution: true -->
```python
# 2.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ANSWER 

def load_dataset(file_name, target_column,fictive_attr_name):
    df = pd.read_csv(file_name)
    y = df[target_column]
    x = df.drop(labels = target_column, axis=1)
    x.insert(loc = 0, column = fictive_attr_name, value=1)
    return x, y
```

```python
# 2.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: SELF TESTING
# Use the following code to test your implementation:
file_name = 'data' + os.sep + 'car_prices_train.csv'
student_name, student_id = myName(), myId()
assignment_params = get_assignment_params(student_id)
# ~~~~ CALL YOUR 'load_dataset' IMPLEMENTATION: ~~~~ #
X, y = load_dataset(file_name, assignment_params['y_col'], assignment_params['fictive_attr'])
# --- add additional code to check your code if needed:
# YOUR CODE HERE

```

```python
# 2.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation, if used)
# === CODE TYPE: SELF TESTING
# ---- Add your additional tests here if needed:
# YOUR CODE HERE

```

<!-- grade_id: test2-1_load_dataset, points: 0.1, grade: true, solution: false -->
```python
# 2.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Parts 2. - Test 1 (name: test2-1_load_dataset, points: 0.1) - Sanity")
print ("\t--->Testing the implementation of 'load_dataset' ...")

file_name = 'data' + os.sep + 'car_prices_train.csv'
try:    
    student_name, student_id = myName(), myId()
    assert type(student_name) is str and type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)
    X, y = load_dataset(file_name, assignment_params['y_col'], assignment_params['fictive_attr'])
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

print ("\nGood Job!\nYou've passed the 1st test for the 'load_dataset' function implementation :-)")
```

Output:
```
Parts 2. - Test 1 (name: test2-1_load_dataset, points: 0.1) - Sanity
	--->Testing the implementation of 'load_dataset' ...

Good Job!
You've passed the 1st test for the 'load_dataset' function implementation :-)
```

<!-- grade_id: test2-2_load_dataset, points: 0.4, grade: true, solution: false -->
```python
# 2.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Parts 2. - Test 2 (name: test2-2_load_dataset, points: 0.4)")
print ("\t--->Testing the implementation of 'load_dataset' ...")

file_name = 'data' + os.sep + 'car_prices_train.csv'
try:    
    student_name, student_id = myName(), myId()
    assert type(student_name) is str and type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)
    X, y = load_dataset(file_name, assignment_params['y_col'], assignment_params['fictive_attr'])
except Exception as e:
    print ('You probably did not implement student-info functions, \nerror Message:',str(e))
    raise

assert assignment_params['fictive_attr'] in X.columns, "Missing fictive column: %r" %(assignment_params['fictive_attr'])
np.testing.assert_array_equal(X.index, y.index, 'X and y should have the same index')
assert (5860, 11) == X.shape, "Wrong shape for feature vectors dataframe"

print ("\nGood Job!\nYou've passed the 2nd test for the 'load_dataset' function implementation :-)")
```

Output:
```
Parts 2. - Test 2 (name: test2-2_load_dataset, points: 0.4)
	--->Testing the implementation of 'load_dataset' ...

Good Job!
You've passed the 2nd test for the 'load_dataset' function implementation :-)
```

<!-- grade_id: test2-3_load_dataset, points: 0.5, grade: true, solution: false -->
```python
# 2.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Parts 2. - Test 3 (name: test2-3_load_dataset, points: 0.5)")
print ("\t--->Testing the implementation of 'load_dataset' ...")

file_name = 'data' + os.sep + 'car_prices_test.csv'
try:    
    student_name, student_id = myName(), myId()
    assert type(student_name) is str and type(student_id) is int, "name is not a string or id is not an integer"  
    assignment_params = get_assignment_params(student_id)
    X, y = load_dataset(file_name, assignment_params['y_col'], assignment_params['fictive_attr'])
except Exception as e:
    print ('You probably did not implement student-info functions, \nerror Message:',str(e))
    raise  

np.testing.assert_array_equal(X.columns, assignment_params['cols'], 'Missmatch in column names or order')
assert X['seats'].head().mean() == 5.4, "Wrong 'seats' column values"
assert y.head().mean()//1000 == 572, "Wrong target values"

print ("\nGood Job!\nYou've passed the 3rd test for the 'load_dataset' function implementation :-)")
```

Output:
```
Parts 2. - Test 3 (name: test2-3_load_dataset, points: 0.5)
	--->Testing the implementation of 'load_dataset' ...

Good Job!
You've passed the 3rd test for the 'load_dataset' function implementation :-)
```

[Go to the beginning of the assignment](#home)
<a id="scaling"></a>
## Part 3 -  scaling 
<img src="./images/architect-scale.jpg" alt="Scale" style="width: 400px;" align='left'/>

<b>Scaling (scale_fit_transform, scale_transform_for_test) - Permitted method usage</b>:<br />
<div class="alert alert-success">
<ul><li>&#x2714; Any function, which calculates the minimum and the maximum (such as numpy's function)</li>
<li>&#x2714; Any function, which calculates the average (such as series's function)</li>
<li>&#x2714; Any function, which calculates a standard deviation/ std/ stddev (such as dataframe's function)
    <ul><li><u>Make sure you use the sample's std</u></li></ul>
<li>&#x2714; Any function, which calculates a variance / var (such as statistic's function)</li>
    <ul><li><u>Make sure you use the sample's variance</u></li></ul>
<li>&#x2714; Any other basic mathematical function (+,*, ...) </li></ul>
</div>
<hr>
<b>Scaling (scale_fit_transform, scale_transform_for_test) - Illegal method usage</b>:<br/>
<div class="alert alert-danger">
    <ul><li>&#x2718; Any function, which performs the scaling</li></ul>
</div>

### Part 3.a. Instructions
<u>Part 3.a.</u> - Scaling (scale dataframe columns, including target);
<u>method name</u>: <b>scale_fit_transform</b><br/>

<pre>
Complete the following function to return a <b>copy</b> of the train set dataframe, 
       with scaled features and scaled target as well as a dataframe containing the scaled trained information.
* Important note 1: The train-set includes both the feature vector and the target, which both should be scaled.
* Important note 2: if you need to implement minmax it should be scaled between -1 to 1.
* Important note 3: The train-set feature vector includes also the fictive feature called 'x_0'.
                    You are expected not to scale this fictive feature
------------
input parameter:
- train_set - the input dataframe, including both the feature vector and the target
- fictive_attr_name - the name of the fictive feature column, with the constant value of 1, 
                      representing the x_0 fictive feature.
- target_col_name   - the name of the column representing the target.                      
------------
return values:
- trained_scaling_info - a dataframe containing the trained info for each column, including the target 
                         (see the example ahead)          
- train_set_scaled - a <b>copy</b> of the input train_set dataframe, with scaled columns, including the target 
                    (the scaling is dependent on the above 'scale_type', which was assigned for you above).
------------
returned statement should look as following:
return trained_scaling_info,  train_set_scaled
---------------------
</pre>
For example, let's assume the value of the input parameter `fictive_attr_name='x_0'`.<br />
Let's assume the value of the input parameter `target_col_name='target'`.<br />

-----

Let's assume the value of the following input `train_set` (including target):

| | x_0 | size | #bedrooms | target |
| :- | :- | :- | :- | :- |
| 0 | 1 | 2104 | 5 | 460 |
| 1 | 1 | 1416 | 3 | 232 |
| 2 | 1 | 852 | 2 | 178 |

------

<u>Case 1</u> - the <b>'scale_type'</b> (which is assigned for you ahead equals) <b>'minmax' [-1,1]</b> - <br />
<b>Important note: minmax here should be scaled between -1 to 1 here</b>
The `trained_scaling_info` output dataframe should look like this:

| | min | max |
| :- | :- | :- |
| size | 852 | 2104 |
| #bedrooms | 2 | 5 |
| target | 178 | 460 |

And the corresponding `train_set_scaled` output dataframe should look like this:

| | x_0 | size | #bedrooms | target |
| :- | :- | :- | :- | :- |
| 0 | 1 | 1 | 1 | 1 |
| 1 | 1 | -0.099041534 | -0.333333333 | -0.617021277 |
| 2 | 1 | -1 | -1 | -1 |

------

<u>Case 2</u> - the <b>scale_type</b> </b> (which is assigned for you ahead equals) <b>'standard'</b> - <br />
The `trained_scaling_info` output dataframe should look like this:

| | mean | std |
| :- | :- | :- |
| Size | 1457.333333 | 627.022594 |
| #bedrooms | 3.333333333 | 1.527525232 |
| target | 290 | 149.6796579 |


And the corresponding `train_set_scaled` output dataframe should look like this:

| | x_0 | size | #bedrooms | target |
| :- | :- | :- | :- | :- |
| 0 | 1 | 1.031329131 | 1.091089451 | 1.135758876  |
| 1 | 1 | -0.065920006 | -0.21821789 | -0.387494205  |
| 2 | 1 | -0.965409124 | -0.872871561 | -0.748264671  |

-----

<b>The return statement should look like this</b>:<br />
return trained_scaling_info, train_set_scaled

### 3. Scaling - you need to <u>implement the following method</u> ...

```python
## 3.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# === CELL TYPE: IMPLEMENTATION TASKS

try:    
    student_name, student_id = myName(), myId()
    assert type(student_name) is str and type(student_id) is int, "name is not a string or id is not an integer"    
    assignment_params = get_assignment_params(student_id)
except Exception as e:
    print ('You probably did not implement student-info functions, \nerror Message:',str(e))
    raise
# ----------------------------
                    
print ("The scaling type you need to implement is: " + assignment_params['scale_type'])
```

Output:
```
The scaling type you need to implement is: t-distribution standardization
```

<!-- grade_id: cell-65aa3333757a1464, points: 0, grade: false, solution: true -->
```python
# 3.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ASSISTANCE TO ANSWER 
# ---- Add assistance code here IF NEEDED:
# YOUR CODE HERE

```

<!-- grade_id: answer-scale_fit_transform, points: 0, grade: false, solution: true -->
```python
# 3.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ANSWER 

def scale_fit_transform(train_set, fictive_attr_name):
    train_set_scaled= train_set.copy()
    fictive_arr = train_set_scaled[fictive_attr_name]
    train_set_scaled = train_set_scaled.drop(labels = fictive_attr_name, axis=1) #without fictive_attr_name
    
    col_mean = train_set_scaled.mean(axis = 0 )
    col_std = train_set_scaled.std(axis = 0 )
    
    trained_scaling_info = pd.DataFrame({'mean':col_mean, 'std':col_std}, index = train_set_scaled.columns)
    
    colm = list(train_set_scaled.columns.values)
    for i in range(len(colm)):
        train_set_scaled[colm[i]] = train_set_scaled[colm[i]].apply(lambda x: (x-col_mean[i])/col_std[i])
    train_set_scaled.insert(loc = 0, column=fictive_attr_name, value=fictive_arr)
    return trained_scaling_info, train_set_scaled
```

```python
# 3.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: SELF TESTING
# Use the following code to test your implementation:
file_name_train = 'data' + os.sep + 'car_prices_train.csv'
student_name, student_id = myName(), myId()
assignment_params = get_assignment_params(student_id)
X_train, y_train = load_dataset(file_name_train, assignment_params['y_col'], assignment_params['fictive_attr'])
y_train.name='target'
train_set=pd.concat([X_train,y_train], axis=1)
trained_scaling_info,  train_set_scaled = scale_fit_transform(train_set, assignment_params['fictive_attr'])
# ~~~~ CALL YOUR 'train_set_scaled' IMPLEMENTATION: ~~~~ #
X_train_scaled,y_train_scaled=train_set_scaled.drop(columns=['target']), train_set_scaled['target']
# --- add additional code to check your code if needed:
# YOUR CODE HERE

```

```python
# 3.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation, if used)
# === CODE TYPE: SELF TESTING
# ---- Add your additional tests here if needed:
# YOUR CODE HERE

```

<!-- grade_id: test3a-1_scale_fit_transform, points: 0.1, grade: true, solution: false -->
```python
# 3.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Parts 3.a. - Test 1 (name: test3a-1_scale_fit_transform, points: 0.1) - Sanity")
print ("\t--->Testing the implementation of 'scale_fit_transform' ...")

file_name_train = 'data' + os.sep + 'car_prices_train.csv'
try:        
    student_name, student_id = myName(), myId()
    assert type(student_name) is str and type(student_id) is int, "name is not a string or id is not an integer"  
    assignment_params = get_assignment_params(student_id)
    print ("\t\tTesing a '" + assignment_params['scale_type'] + "' scale type ...")
    X_train, y_train = load_dataset(file_name_train, assignment_params['y_col'], assignment_params['fictive_attr'])
    y_train.name='target'
    train_set=pd.concat([X_train,y_train], axis=1)
    trained_scaling_info,  train_set_scaled = scale_fit_transform(train_set, assignment_params['fictive_attr'])
    X_train_scaled,y_train_scaled=train_set_scaled.drop(columns=['target']), train_set_scaled['target']
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

print ("\nGood Job!\nYou've passed the 1st test for the 'scale_fit_transform' function implementation :-)")
```

Output:
```
Parts 3.a. - Test 1 (name: test3a-1_scale_fit_transform, points: 0.1) - Sanity
	--->Testing the implementation of 'scale_fit_transform' ...
		Tesing a 't-distribution standardization' scale type ...

Good Job!
You've passed the 1st test for the 'scale_fit_transform' function implementation :-)
```

<!-- grade_id: test3a-2_scale_fit_transform, points: 0.1, grade: true, solution: false -->
```python
# 3.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Parts 3.a. - Test 2 (name: test3a-2_scale_fit_transform, points: 0.1)")
print ("\t--->Testing the implementation of 'scale_fit_transform' ...")

file_name_train = 'data' + os.sep + 'car_prices_train.csv'
try:        
    student_name, student_id = myName(), myId()
    assert type(student_name) is str and type(student_id) is int, "name is not a string or id is not an integer"  
    assignment_params = get_assignment_params(student_id)
    print ("\t\tTesing a '" + assignment_params['scale_type'] + "' scale type ...")
    X_train, y_train = load_dataset(file_name_train, assignment_params['y_col'], assignment_params['fictive_attr'])
    y_train.name='target'
    train_set=pd.concat([X_train,y_train], axis=1)
    trained_scaling_info,  train_set_scaled = scale_fit_transform(train_set, assignment_params['fictive_attr'])
    X_train_scaled,y_train_scaled=train_set_scaled.drop(columns=['target']), train_set_scaled['target']
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
np.testing.assert_array_equal(train_set_scaled.columns, assignment_params['cols']+['target'], 'scaled test and test dataframes should have the same column names')

print ("\nGood Job!\nYou've passed the 2nd test for the 'scale_fit_transform' function implementation :-)")
```

Output:
```
Parts 3.a. - Test 2 (name: test3a-2_scale_fit_transform, points: 0.1)
	--->Testing the implementation of 'scale_fit_transform' ...
		Tesing a 't-distribution standardization' scale type ...

Good Job!
You've passed the 2nd test for the 'scale_fit_transform' function implementation :-)
```

<!-- grade_id: test3a-3_scale_fit_transform, points: 0.5, grade: true, solution: false -->
```python
# 3.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Parts 3.a. - Test 3 (name: test3a-3_scale_fit_transform, points: 0.5)")
print ("\t--->Testing the implementation of 'scale_fit_transform' ...")

file_name_train = 'data' + os.sep + 'car_prices_train.csv'
try:        
    student_name, student_id = myName(), myId()
    assert type(student_name) is str and type(student_id) is int, "name is not a string or id is not an integer"  
    assignment_params = get_assignment_params(student_id)
    print ("\t\tTesing a '" + assignment_params['scale_type'] + "' scale type ...")
    X_train, y_train = load_dataset(file_name_train, assignment_params['y_col'], assignment_params['fictive_attr'])
    y_train.name='target'
    train_set=pd.concat([X_train,y_train], axis=1)
    trained_scaling_info,  train_set_scaled = scale_fit_transform(train_set, assignment_params['fictive_attr'])
    X_train_scaled,y_train_scaled=train_set_scaled.drop(columns=['target']), train_set_scaled['target']
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

if 't-distribution standardization' == assignment_params['scale_type']:
    assert np.round(train_set_scaled.max().max(),2) == 38.46, 'Wrong max standardized scaled values'
else:    
    assert np.round(train_set_scaled.mean().min()+0.0006,2) == -0.94, 'Wrong average minmax scaled values'
    
print ("\nGood Job!\nYou've passed the 3rd test for the 'scale_fit_transform' function implementation :-)")
```

Output:
```
Parts 3.a. - Test 3 (name: test3a-3_scale_fit_transform, points: 0.5)
	--->Testing the implementation of 'scale_fit_transform' ...
		Tesing a 't-distribution standardization' scale type ...

Good Job!
You've passed the 3rd test for the 'scale_fit_transform' function implementation :-)
```

### Part 3.b. Instructions
<u>Part 3.b.</u> - Scaling (scale dataframe columns, including target);
<u>method name</u>: <b>scale_transform_for_test</b><br/>
<pre>
Complete the following function to return a <b>copy</b> of the test-set dataframe, 
       with scaled features and scaled target.
* Important note 1: The test-set includes both the feature vector and the target, which both should be scaled.
* Important note 2: if you need to implement minmax it should be scaled between -1 to 1.       
* Important note 3: The train-set feature vector includes also the fictive feature called 'x_0'.
                    You are expected not to scale this fictive feature
------------
input parameters:
- trained_scaling_info - a dataframe containing the trained info for each column, including the target 
                         (see the example ahead)    
- test_set - the input test-set dataframe, including both the feature vector and the target
- fictive_attr_name - the name of the fictive feature column, with the constant value of 1, 
                      representing the x_0 fictive feature.
------------
return values:
- test_set_scaled - a <b>copy</b> of the input train dataframe, with scaled columns, including the target 
                   (the scaling is dependent on the above 'scale_type', which was assigned for you above)
------------
returned statement should look as following:
return test_set_scaled
---------------------
</pre>
For example, let's assume the following values for the input parameters:
* `fictive_attr_name='x_0'`
-----

The `trained_scaling_info` is from the output of the `scale_fit_transform` previous method.<br />
The corresponding input possibilities, depend on the implementation, assigned for you.<br />

<u>Expected input case 1</u> - the <b>'scale_type'='minmax' [-1,1]</b><br />
<b>Important note: minmax here should be scaled between -1 to 1 here</b><br />
Let's assume we've got the following input trained_scaling_info parameter (including target):

| | min | max |
| :- | :- | :- |
| size | 852 | 2104 |
| #bedrooms | 2 | 5 |
| target | 178 | 460 |

-----

<u>Expected input case 2</u> - the <b>'scale_type'='standard'</b><br />
Let's assume we've got the following input `trained_scaling_info` parameter (including target):

| | mean | std |
| :- | :- | :- |
| Size | 1457.333333 | 627.022594 |
| #bedrooms | 3.333333333 | 1.527525232 |
| target | 290 | 149.6796579 |

-----

Let's assume we've got the following input `test_set` parameter (including target):

| | x_0 | size | #bedrooms | target |
| :- | :- | :- | :- | :- |
| 0 | 1 | 1022 | 2 | 200 |
| 1 | 1 | 1300 | 2.5 | 220 |
| 2 | 1 | 1900 | 5 | 410 |

------

<u>Expected output case 1</u> - the <b>'scale_type'</b> (which is assigned for you ahead equals) <b>'minmax' [-1,1]</b> - <br />

The corresponding `test_set_scaled` output dataframe should look like this:

| | x_0 | size | #bedrooms | target |
| :- | :- | :- | :- | :- |
| 0 | 1 | -0.728434505 | -1 | -0.843971631 |
| 1 | 1 | -0.284345048 | -0.666666667 | -0.70212766 |
| 2 | 1 | 0.674121406 | 1 | 0.645390071 |

------

<u>Expected output case 2</u> - the <b>scale_type</b> </b> (which is assigned for you ahead equals) <b>'standard'</b> - <br />

The corresponding `test_set_scaled` output dataframe should look like this:

| | x_0 | size | #bedrooms | target |
| :- | :- | :- | :- | :- |
| 0 | 1 | -0.694286518 | -0.872871561 | -0.601284111 |
| 1 | 1 | -0.250921314 | -0.545544726 | -0.467665419 |
| 2 | 1 | 0.705982003 | 1.091089451 | 0.801712147 |

-----

<b>The return statement should look like this</b>:<br />
return test_set_scaled

<!-- grade_id: cell-032134bb90236b88, points: 0, grade: false, solution: true -->
```python
# 3.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ASSISTANCE TO ANSWER 
# ---- Add assistance code here IF NEEDED:
# YOUR CODE HERE

```

<!-- grade_id: answer-scale_transform_for_test, points: 0, grade: false, solution: true -->
```python
# 3.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ANSWER 

def scale_transform_for_test(trained_scaling_info,  test_set, fictive_attr_name):
    test_set_scaled = test_set.copy()
    fictive_arr = test_set_scaled[fictive_attr_name];
    test_set_scaled = test_set_scaled.drop(labels = fictive_attr_name, axis=1) #without fictive_attr_name
    
    col_mean = trained_scaling_info['mean'].values
    col_std = trained_scaling_info['std'].values
    colm = list(test_set.columns.values)
    colm = list(test_set_scaled.columns.values)
    for i in range(len(colm)):
        test_set_scaled[colm[i]] = test_set_scaled[colm[i]].apply(lambda x: (x-col_mean[i])/col_std[i])
    test_set_scaled.insert(loc = 0, column=fictive_attr_name, value=fictive_arr)
    
    return test_set_scaled
```

```python
# 3.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: SELF TESTING
# Use the following code to test your implementation:
# ---- load and scale train-set:
file_name_train = 'data' + os.sep + 'car_prices_train.csv'
student_name, student_id = myName(), myId()
assignment_params = get_assignment_params(student_id)
X_train, y_train = load_dataset(file_name_train, assignment_params['y_col'], assignment_params['fictive_attr'])
y_train.name='target'
train_set=pd.concat([X_train,y_train], axis=1)
trained_scaling_info,  train_set_scaled = scale_fit_transform(train_set, assignment_params['fictive_attr'])
X_train_scaled,y_train_scaled=train_set_scaled.drop(columns=['target']), train_set_scaled['target']
# ---- load and scale test-set:
file_name_test  = 'data' + os.sep + 'car_prices_test.csv'
X_test, y_test   = load_dataset(file_name_test, assignment_params['y_col'], assignment_params['fictive_attr'])
y_test.name='target'
test_set =pd.concat([X_test,y_test], axis=1)   
# ~~~~ CALL YOUR 'scale_transform_for_test' IMPLEMENTATION: ~~~~ #
test_set_scaled = scale_transform_for_test(trained_scaling_info,  test_set, assignment_params['fictive_attr'])
X_test_scaled,y_test_scaled=test_set_scaled.drop(columns=['target']), test_set_scaled['target']
# --- add additional code to check your code if needed:
# YOUR CODE HERE

```

```python
# 3.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation, if used)
# === CODE TYPE: SELF TESTING
# ---- Add your additional tests here if needed:
# YOUR CODE HERE

```

<!-- grade_id: test3b-1_scale_transform_for_test, points: 0.1, grade: true, solution: false -->
```python
# 3.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Parts 3.b. - Test 1 (name: test3b-1_scale_transform_for_test, points: 0.1) - Sanity")
print ("\t--->Testing the implementation of 'scale_transform_for_test' ...")

file_name_train = 'data' + os.sep + 'car_prices_train.csv'
file_name_test  = 'data' + os.sep + 'car_prices_test.csv'
try:        
    student_name, student_id = myName(), myId()
    assert type(student_name) is str and type(student_id) is int, "name is not a string or id is not an integer"  
    assignment_params = get_assignment_params(student_id)
    X_train, y_train = load_dataset(file_name_train, assignment_params['y_col'], assignment_params['fictive_attr'])
    y_train.name='target'
    train_set=pd.concat([X_train,y_train], axis=1)
    trained_scaling_info,  train_set_scaled = scale_fit_transform(train_set, assignment_params['fictive_attr'])
    X_train_scaled,y_train_scaled=train_set_scaled.drop(columns=['target']), train_set_scaled['target']
    file_name_test  = 'data' + os.sep + 'car_prices_test.csv'
    X_test, y_test   = load_dataset(file_name_test, assignment_params['y_col'], assignment_params['fictive_attr'])
    y_test.name='target'
    test_set =pd.concat([X_test,y_test], axis=1)   
    test_set_scaled = scale_transform_for_test(trained_scaling_info, test_set, assignment_params['fictive_attr'])
    X_test_scaled,y_test_scaled=test_set_scaled.drop(columns=['target']), test_set_scaled['target']
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

print ("\nGood Job!\nYou've passed the 1st test for the 'scale_transform_for_test' function implementation :-)")
```

Output:
```
Parts 3.b. - Test 1 (name: test3b-1_scale_transform_for_test, points: 0.1) - Sanity
	--->Testing the implementation of 'scale_transform_for_test' ...

Good Job!
You've passed the 1st test for the 'scale_transform_for_test' function implementation :-)
```

<!-- grade_id: test3b-2_scale_transform_for_test, points: 0.1, grade: true, solution: false -->
```python
# 3.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Parts 3.b. - Test 2 (name: test3b-2_scale_transform_for_test, points: 0.1)")
print ("\t--->Testing the implementation of 'scale_transform_for_test' ...")

file_name_train = 'data' + os.sep + 'car_prices_train.csv'
file_name_test  = 'data' + os.sep + 'car_prices_test.csv'
try:        
    student_name, student_id = myName(), myId()
    assert type(student_name) is str and type(student_id) is int, "name is not a string or id is not an integer"  
    assignment_params = get_assignment_params(student_id)
    print ("\t\tTesing a '" + assignment_params['scale_type'] + "' scale type ...")
    X_train, y_train = load_dataset(file_name_train, assignment_params['y_col'], assignment_params['fictive_attr'])
    X_test, y_test   = load_dataset(file_name_test, assignment_params['y_col'], assignment_params['fictive_attr'])
    y_train.name,y_test.name='target','target'
    train_set=pd.concat([X_train,y_train], axis=1)
    test_set =pd.concat([X_test,y_test], axis=1)   
    trained_scaling_info,  train_set_scaled = scale_fit_transform(train_set, assignment_params['fictive_attr'])
    test_set_scaled = scale_transform_for_test(trained_scaling_info,  test_set, assignment_params['fictive_attr'])
    X_train_scaled,y_train_scaled=train_set_scaled.drop(columns=['target']), train_set_scaled['target']
    X_test_scaled,y_test_scaled=test_set_scaled.drop(columns=['target']), test_set_scaled['target']
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
np.testing.assert_array_equal(train_set_scaled.columns, test_set_scaled.columns, 'scaled train and test dataframes should have the same column names')

print ("\nGood Job!\nYou've passed the 2nd test for the 'scale_transform_for_test' function implementation :-)")
```

Output:
```
Parts 3.b. - Test 2 (name: test3b-2_scale_transform_for_test, points: 0.1)
	--->Testing the implementation of 'scale_transform_for_test' ...
		Tesing a 't-distribution standardization' scale type ...

Good Job!
You've passed the 2nd test for the 'scale_transform_for_test' function implementation :-)
```

<!-- grade_id: test3b-3_scale_transform_for_test, points: 0.6, grade: true, solution: false -->
```python
# 3.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Parts 3.b. - Test 3 (name: test3b-3_scale_transform_for_test, points: 0.6)")
print ("\t--->Testing the implementation of 'scale_transform_for_test' ...")

file_name_train = 'data' + os.sep + 'car_prices_train.csv'
file_name_test  = 'data' + os.sep + 'car_prices_test.csv'
try:        
    student_name, student_id = myName(), myId()
    assert type(student_name) is str and type(student_id) is int, "name is not a string or id is not an integer"  
    assignment_params = get_assignment_params(student_id)
    print ("\t\tTesing a '" + assignment_params['scale_type'] + "' scale type ...")
    X_train, y_train = load_dataset(file_name_train, assignment_params['y_col'], assignment_params['fictive_attr'])
    X_test, y_test   = load_dataset(file_name_test, assignment_params['y_col'], assignment_params['fictive_attr'])
    y_train.name,y_test.name='target','target'
    train_set=pd.concat([X_train,y_train], axis=1)
    test_set =pd.concat([X_test,y_test], axis=1)   
    
    trained_scaling_info,  train_set_scaled = scale_fit_transform(train_set, assignment_params['fictive_attr'])
    test_set_scaled = scale_transform_for_test(trained_scaling_info,  test_set, assignment_params['fictive_attr'])
    X_train_scaled,y_train_scaled=train_set_scaled.drop(columns=['target']), train_set_scaled['target']
    X_test_scaled,y_test_scaled=test_set_scaled.drop(columns=['target']), test_set_scaled['target']
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
if 't-distribution standardization' == assignment_params['scale_type']:
    assert np.round(test_set_scaled.max().mean()-0.0005,2)== 3.39, 'Wrong max standardized scaled values'
else:
    assert np.round(test_set_scaled.min().mean()+0.005,2) == -0.8, 'Wrong min minmax scaled values'
    
print ("\nGood Job!\nYou've passed the 3rd test for the 'scale_transform_for_test' functions implementation :-)")
```

Output:
```
Parts 3.b. - Test 3 (name: test3b-3_scale_transform_for_test, points: 0.6)
	--->Testing the implementation of 'scale_transform_for_test' ...
		Tesing a 't-distribution standardization' scale type ...

Good Job!
You've passed the 3rd test for the 'scale_transform_for_test' functions implementation :-)
```

[Go to the beginning of the assignment](#home)
<a id="predict"></a>
### 4. Predict

### Part 4. Instructions
<u>Part 4.</u> - predict</u> (Linear regression - predict); 
<u>method name</u>: <b>predict</b><br/>

<pre>
Change the following method to calculate and return the linear regression prediction 
(for either the multivariate or the univariate case).
 Note: the 'predict' method should be used both for the predictions on the test-set,
       as well as during the training process.     
------------
input parameters:
- x_feature_vectors - a dataframe containing all feature vectors for which we want to predict the y_predicted values,
                      using the linear regression model's weights.
                 Note: it is expected the the x_feature_vectors contain a fictive column of features, with 
                       the constant value of 1, representing the x_0 fictive feature (called 'x_0').
- trained_w - the trained w-vector, i.e. the trained weight vector 
              (or the weight vector from some step in the training).
                 Note: it is expected the the trained_w series contain, the w_0 cell as a first cell.
                       The index of all cells are expected to fit the 'columns' (i.e. column titles which
                       are the column indexes) of the 'x_feature_vectors' dataframe, including the 'x_0' 
                       fictive column.
------------
return value:
- y_predicted - a series of the predictions for each input feature vector instance
------------
The return statement should look like the following:
return y_predicted
</pre>

<!-- grade_id: cell-58f61e66de055e12, points: 0, grade: false, solution: true -->
```python
# 4.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ASSISTANCE TO ANSWER 
# ---- Add assistance code here IF NEEDED:
# YOUR CODE HERE

```

<!-- grade_id: answer-predict, points: 0, grade: false, solution: true -->
```python
# 4.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ANSWER 

def predict(x_featureVectors,trained_w):
    x_featureMul = x_featureVectors.copy()
    colm = x_featureMul.columns.values
    for i in range(len(colm)):
        x_featureMul[colm[i]] = x_featureMul[colm[i]].apply(lambda x: (x* trained_w[i]))
    y_predicted = x_featureMul.sum(axis = 1);
    print(y_predicted)
    return y_predicted
```

```python
# 4.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: SELF TESTING
# Use the following code to test your implementation:
# ---- load and scale test-set:
file_name_test  = 'data' + os.sep + 'car_prices_test.csv'
X_test, y_test   = load_dataset(file_name_test, assignment_params['y_col'], assignment_params['fictive_attr'])
# ---- predict (some of) the test examples (with some of the features), 
#              using a fake w vector:
w_fake = pd.Series([2.5, 0.1],index=X_train_scaled.iloc[:,:2].columns)
# ~~~~ CALL YOUR 'predict' IMPLEMENTATION: ~~~~ #
y_predicted = predict(X_test.iloc[:100,:2],w_fake)
# --- add additional code to check your code if needed:
# YOUR CODE HERE

```

Output:
```
0     204.2
1     204.2
2     203.8
3     204.2
4     203.5
      ...  
95    204.0
96    203.8
97    203.7
98    203.9
99    203.7
Length: 100, dtype: float64
```

```python
# 4.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation, if used)
# === CODE TYPE: SELF TESTING
# ---- Add your additional tests here if needed:
# YOUR CODE HERE

```

<!-- grade_id: test4-1_predict, points: 0.1, grade: true, solution: false -->
```python
# 4.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Parts 4. - Test 1 (name: test4-1_predict, points: 0.1) - Sanity")
print ("\t--->Testing the implementation of 'predict' ...")

file_name_train = 'data' + os.sep + 'car_prices_train.csv'
file_name_test  = 'data' + os.sep + 'car_prices_test.csv'
try:        
    student_name, student_id = myName(), myId()
    assert type(student_name) is str and type(student_id) is int, "name is not a string or id is not an integer"  
    assignment_params = get_assignment_params(student_id)
    X_train, y_train = load_dataset(file_name_train, assignment_params['y_col'], assignment_params['fictive_attr'])
    X_test, y_test   = load_dataset(file_name_test, assignment_params['y_col'], assignment_params['fictive_attr'])
    y_train.name,y_test.name='target','target'
    train_set=pd.concat([X_train,y_train], axis=1)
    test_set =pd.concat([X_test,y_test], axis=1)   
    trained_scaling_info,  train_set_scaled = scale_fit_transform(train_set, assignment_params['fictive_attr'])
    test_set_scaled = scale_transform_for_test(trained_scaling_info,  test_set, assignment_params['fictive_attr'])
    X_train_scaled,y_train_scaled=train_set_scaled.drop(columns=['target']), train_set_scaled['target']
    X_test_scaled,y_test_scaled=test_set_scaled.drop(columns=['target']), test_set_scaled['target']
    w_fake = pd.Series([2.5, 0.1],index=X_train_scaled.iloc[:,:2].columns)
    y_predicted = predict(X_test.iloc[:100,:2],w_fake)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

print ("\nGood Job!\nYou've passed the 1st test for the 'predict' function implementation :-)")
```

Output:
```
Parts 4. - Test 1 (name: test4-1_predict, points: 0.1) - Sanity
	--->Testing the implementation of 'predict' ...
0     204.2
1     204.2
2     203.8
3     204.2
4     203.5
      ...  
95    204.0
96    203.8
97    203.7
98    203.9
99    203.7
Length: 100, dtype: float64

Good Job!
You've passed the 1st test for the 'predict' function implementation :-)
```

<!-- grade_id: test4-2_predict, points: 0.4, grade: true, solution: false -->
```python
# 4.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Parts 4. - Test 2 (name: test4-2_predict, points: 0.4)")
print ("\t--->Testing the implementation of 'predict' ...")

file_name_train = 'data' + os.sep + 'car_prices_train.csv'
file_name_test  = 'data' + os.sep + 'car_prices_test.csv'
try:        
    student_name, student_id = myName(), myId()
    assert type(student_name) is str and type(student_id) is int, "name is not a string or id is not an integer"  
    assignment_params = get_assignment_params(student_id)
    X_train, y_train = load_dataset(file_name_train, assignment_params['y_col'], assignment_params['fictive_attr'])
    X_test, y_test   = load_dataset(file_name_test, assignment_params['y_col'], assignment_params['fictive_attr'])
    y_train.name,y_test.name='target','target'
    train_set=pd.concat([X_train,y_train], axis=1)
    test_set =pd.concat([X_test,y_test], axis=1)   
    trained_scaling_info,  train_set_scaled = scale_fit_transform(train_set, assignment_params['fictive_attr'])
    test_set_scaled = scale_transform_for_test(trained_scaling_info,  test_set, assignment_params['fictive_attr'])
    X_train_scaled,y_train_scaled=train_set_scaled.drop(columns=['target']), train_set_scaled['target']
    X_test_scaled,y_test_scaled=test_set_scaled.drop(columns=['target']), test_set_scaled['target']
    w_fake = pd.Series([2.5, 0.1],index=X_train_scaled.iloc[:,:2].columns)
    y_predicted = predict(X_test_scaled.iloc[5:7,:2],w_fake)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

if 't-distribution standardization' == assignment_params['scale_type']:
    assert 2.45 == np.round(y_predicted.values[0],2), 'Wrong output for predict'
else:    
    assert 2.56 == np.round(y_predicted.values[1]-0.0005,2), 'Wrong output for predict'

print ("\nGood Job!\nYou've passed the 2nd test for the 'predict' function implementation :-)")
```

Output:
```
Parts 4. - Test 2 (name: test4-2_predict, points: 0.4)
	--->Testing the implementation of 'predict' ...
5    2.450072
6    2.527480
dtype: float64

Good Job!
You've passed the 2nd test for the 'predict' function implementation :-)
```

<!-- grade_id: test4-3_predict, points: 0.5, grade: true, solution: false -->
```python
# 4.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Parts 4. - Test 3 (name: test4-3_predict, points: 0.5)")
print ("\t--->Testing the implementation of 'predict' ...")
print ("\n\t====> Full grading test - the following test is invisible. You will not know if you passed it, before grades are returned to the students")

```

Output:
```
Parts 4. - Test 3 (name: test4-3_predict, points: 0.5)
	--->Testing the implementation of 'predict' ...

	====> Full grading test - the following test is invisible. You will not know if you passed it, before grades are returned to the students
```

[Go to the beginning of the assignment](#home)
<a id="fit"></a>
### 5. Fit - training
<img src="./images/gradientDescent_animation.gif" alt="Linear Regression w/Gradient Descent" width="400" align='left'/>

### Part 5. Instructions
<u>Part 5.</u> - multivariate linear regression</u> (training with gradient descent);
<u>method name</u>: <b>fit</b><br/>

<pre>
Change the following method to train a multivariate (two or more features) linear regression model, 
using the gradient descent algorithm and returns the trained wight vector w.
  Notes: 
  * note that the initial value of the weight vectors should be 0 for every attribute (including w_0).
  Note also: you should use the predict method for interim prediction, during training.
------------
input parameters:
- X_train - a dataframe containing all feature vectors of the train set.
- y_train - a series of containing all target values per train instance.
- alpha - the learning rate
- num_iterations - number of iterations to run multivariate gradient descent
------------
return value:
- trained_w - a series containing the values of the vector of weights (w), including w_0
              * note the indexes should be equal to the 'X_train' dataframe's column names
                (which are actually column indexes), including 'x_0'
------------
The return statement should look like the following:
return trained_w
</pre>

#### possible output:
The possible 'trained_w' might look like this:

| | |
| :- | :- |
| x_0 | -0.00524 |
| year | -2.33335 |
| is_diesel | -1.11111 |
| owner | 3.00097 |
| is_automatic | -0.11424 |
| km_driven | -1.03333 |
| seats | 2.545454 |
| mileage (kmpl) | 0.121212 |
| max_horse_power | 0.972972 |
| is_individual_seller | -0.012013014 |
| engine_capacity_(cc) | 0.100101102103 |

<!-- grade_id: cell-5facc5dfbf2b358a, points: 0, grade: false, solution: true -->
```python
# 5.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ASSISTANCE TO ANSWER 
# ---- Add assistance code here IF NEEDED:
# YOUR CODE HERE

```

<!-- grade_id: answer-fit, points: 0, grade: false, solution: true -->
```python
# 5.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ANSWER 

def fit(X_train,y_train,alpha,num_iterations):
    from IPython.display import display
    row = list(X_train.columns.values) 
    trained_w = pd.Series(0, index= X_train.columns.values)
    trained_w_temp = trained_w.copy()
    for i in range(num_iterations):
        y_predicted = predict(X_train,trained_w)
        display(y_predicted)
        delta = y_predicted - y_train
        for ii in range(X_train.shape[1]): #num of culomn
            sum1 = 0
            if ii == 0:
                sum1 = delta.sum()
                total = alpha * 2 * (1/(X_train.shape[0])) *sum1
                trained_w_temp.loc[row[0]] = trained_w_temp.loc[row[0]] - total
            else :
                sum1 = 0
                for iii in range(X_train.shape[0]):
                    temp = 0
                    temp = X_train.iloc[iii,ii]
                    temp = temp * delta.iloc[iii]
                    sum1 += temp
                total = alpha * 2 * (1/(X_train.shape[0])) *sum1
                trained_w_temp.loc[row[ii]] =trained_w_temp.loc[row[ii]] - total
        trained_w = trained_w_temp.copy()
    
    display(trained_w)
    return trained_w
```

```python
# 5.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: SELF TESTING
# Use the following code to test your implementation:
# ---- load and scale train-set:
file_name_train = 'data' + os.sep + 'car_prices_train.csv'
student_name, student_id = myName(), myId()
assignment_params = get_assignment_params(student_id)
X_train, y_train = load_dataset(file_name_train, assignment_params['y_col'], assignment_params['fictive_attr'])
y_train.name='target'
train_set=pd.concat([X_train,y_train], axis=1)
trained_scaling_info,  train_set_scaled = scale_fit_transform(train_set, assignment_params['fictive_attr'])
# ---- fit [only 1D (1st 2 columns)] - perform training and get model:
X_train_scaled,y_train_scaled=train_set_scaled.drop(columns=['target']).iloc[:,:2], train_set_scaled['target']
alpha,num_iterations = 0.01,1
# ~~~~ CALL YOUR 'fit' IMPLEMENTATION: ~~~~ #
trained_w = fit(X_train_scaled,y_train_scaled,alpha,num_iterations)    
# --- add additional code to check your code if needed:
# YOUR CODE HERE

```

Output:
```
0       0.0
1       0.0
2       0.0
3       0.0
4       0.0
       ... 
5855    0.0
5856    0.0
5857    0.0
5858    0.0
5859    0.0
Length: 5860, dtype: float64
0       0.0
1       0.0
2       0.0
3       0.0
4       0.0
       ... 
5855    0.0
5856    0.0
5857    0.0
5858    0.0
5859    0.0
Length: 5860, dtype: float64
x_0     9.336483e-19
year    8.157969e-03
dtype: float64
```

```python
# 5.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation, if used)
# === CODE TYPE: SELF TESTING
# ---- Add your additional tests here if needed:
# YOUR CODE HERE

```

<!-- grade_id: test5-1_fit, points: 0.1, grade: true, solution: false -->
```python
# 5.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Parts 5. - Test 1 (name: test5-1_fit, points: 0.1) - Sanity")
print ("\t--->Testing the implementation of 'fit' ...")

file_name_train = 'data' + os.sep + 'car_prices_train.csv'
file_name_test  = 'data' + os.sep + 'car_prices_test.csv'
try:        
    student_name, student_id = myName(), myId()
    assert type(student_name) is str and type(student_id) is int, "name is not a string or id is not an integer"  
    assignment_params = get_assignment_params(student_id)
    X_train, y_train = load_dataset(file_name_train, assignment_params['y_col'], assignment_params['fictive_attr'])
    X_test, y_test   = load_dataset(file_name_test, assignment_params['y_col'], assignment_params['fictive_attr'])
    y_train.name,y_test.name='target','target'
    train_set=pd.concat([X_train,y_train], axis=1)
    test_set =pd.concat([X_test,y_test], axis=1)   
    trained_scaling_info,  train_set_scaled = scale_fit_transform(train_set, assignment_params['fictive_attr'])
    test_set_scaled = scale_transform_for_test(trained_scaling_info,  test_set, assignment_params['fictive_attr'])
    X_train_scaled,y_train_scaled=train_set_scaled.drop(columns=['target']).iloc[:,:2], train_set_scaled['target']
    X_test_scaled,y_test_scaled=test_set_scaled.drop(columns=['target']).iloc[:,:2], test_set_scaled['target']
    alpha,num_iterations = 0.01,1
    trained_w = fit(X_train_scaled,y_train_scaled,alpha,num_iterations)    
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')

print ("\nGood Job!\nYou've passed the 1st test for the 'fit' function implementation :-)")
```

Output:
```
Parts 5. - Test 1 (name: test5-1_fit, points: 0.1) - Sanity
	--->Testing the implementation of 'fit' ...
0       0.0
1       0.0
2       0.0
3       0.0
4       0.0
       ... 
5855    0.0
5856    0.0
5857    0.0
5858    0.0
5859    0.0
Length: 5860, dtype: float64
0       0.0
1       0.0
2       0.0
3       0.0
4       0.0
       ... 
5855    0.0
5856    0.0
5857    0.0
5858    0.0
5859    0.0
Length: 5860, dtype: float64
x_0     9.336483e-19
year    8.157969e-03
dtype: float64

Good Job!
You've passed the 1st test for the 'fit' function implementation :-)
```

<!-- grade_id: test5-2_fit, points: 0.2, grade: true, solution: false -->
```python
# 5.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Parts 5. - Test 2 (name: test5-2_fit, points: 0.2) - Sanity (2)")
print ("\t--->Testing the implementation of 'fit' ...")

file_name_train = 'data' + os.sep + 'car_prices_train.csv'
file_name_test  = 'data' + os.sep + 'car_prices_test.csv'
try:        
    student_name, student_id = myName(), myId()
    assert type(student_name) is str and type(student_id) is int, "name is not a string or id is not an integer"  
    assignment_params = get_assignment_params(student_id)
    X_train, y_train = load_dataset(file_name_train, assignment_params['y_col'], assignment_params['fictive_attr'])
    X_test, y_test   = load_dataset(file_name_test, assignment_params['y_col'], assignment_params['fictive_attr'])
    y_train.name,y_test.name='target','target'
    train_set=pd.concat([X_train,y_train], axis=1)
    test_set =pd.concat([X_test,y_test], axis=1)   
    trained_scaling_info,  train_set_scaled = scale_fit_transform(train_set, assignment_params['fictive_attr'])
    test_set_scaled = scale_transform_for_test(trained_scaling_info,  test_set, assignment_params['fictive_attr'])
    X_train_scaled,y_train_scaled=train_set_scaled.drop(columns=['target']).iloc[:,:2], train_set_scaled['target']
    X_test_scaled,y_test_scaled=test_set_scaled.drop(columns=['target']).iloc[:,:2], test_set_scaled['target']
    alpha,num_iterations = 0.01,2
    trained_w = fit(X_train_scaled,y_train_scaled,alpha,num_iterations)    
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')

print ("\nGood Job!\nYou've passed the 2nd test for the 'fit' function implementation :-)")
```

Output:
```
Parts 5. - Test 2 (name: test5-2_fit, points: 0.2) - Sanity (2)
	--->Testing the implementation of 'fit' ...
0       0.0
1       0.0
2       0.0
3       0.0
4       0.0
       ... 
5855    0.0
5856    0.0
5857    0.0
5858    0.0
5859    0.0
Length: 5860, dtype: float64
0       0.0
1       0.0
2       0.0
3       0.0
4       0.0
       ... 
5855    0.0
5856    0.0
5857    0.0
5858    0.0
5859    0.0
Length: 5860, dtype: float64
0       0.004347
1       0.004347
2      -0.004073
3      -0.006178
4       0.008557
          ...   
5855    0.010662
5856    0.000137
5857    0.000137
5858   -0.006178
5859    0.008557
Length: 5860, dtype: float64
0       0.004347
1       0.004347
2      -0.004073
3      -0.006178
4       0.008557
          ...   
5855    0.010662
5856    0.000137
5857    0.000137
5858   -0.006178
5859    0.008557
Length: 5860, dtype: float64
x_0     1.903673e-18
year    1.615281e-02
dtype: float64

Good Job!
You've passed the 2nd test for the 'fit' function implementation :-)
```

<!-- grade_id: test5-3_fit, points: 0.2, grade: true, solution: false -->
```python
# 5.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Parts 5. - Test 3 (name: test5-3_fit, points: 0.2) - Sanity (3)")
print ("\t--->Testing the implementation of 'fit' ...")

file_name_train = 'data' + os.sep + 'car_prices_train.csv'
file_name_test  = 'data' + os.sep + 'car_prices_test.csv'
try:        
    student_name, student_id = myName(), myId()
    assert type(student_name) is str and type(student_id) is int, "name is not a string or id is not an integer"  
    assignment_params = get_assignment_params(student_id)
    X_train, y_train = load_dataset(file_name_train, assignment_params['y_col'], assignment_params['fictive_attr'])
    X_test, y_test   = load_dataset(file_name_test, assignment_params['y_col'], assignment_params['fictive_attr'])
    y_train.name,y_test.name='target','target'
    train_set=pd.concat([X_train,y_train], axis=1)
    test_set =pd.concat([X_test,y_test], axis=1)   
    trained_scaling_info,  train_set_scaled = scale_fit_transform(train_set, assignment_params['fictive_attr'])
    test_set_scaled = scale_transform_for_test(trained_scaling_info,  test_set, assignment_params['fictive_attr'])
    X_train_scaled,y_train_scaled=train_set_scaled.drop(columns=['target']).iloc[:,:2], train_set_scaled['target']
    X_test_scaled,y_test_scaled=test_set_scaled.drop(columns=['target']).iloc[:,:2], test_set_scaled['target']
    alpha,num_iterations = 0.01,10
    trained_w = fit(X_train_scaled,y_train_scaled,alpha,num_iterations)    
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')

print ("\nGood Job!\nYou've passed the 3rd test for the 'fit' function implementation :-)")
```

Output:
```
Parts 5. - Test 3 (name: test5-3_fit, points: 0.2) - Sanity (3)
	--->Testing the implementation of 'fit' ...
0       0.0
1       0.0
2       0.0
3       0.0
4       0.0
       ... 
5855    0.0
5856    0.0
5857    0.0
5858    0.0
5859    0.0
Length: 5860, dtype: float64
0       0.0
1       0.0
2       0.0
3       0.0
4       0.0
       ... 
5855    0.0
5856    0.0
5857    0.0
5858    0.0
5859    0.0
Length: 5860, dtype: float64
0       0.004347
1       0.004347
2      -0.004073
3      -0.006178
4       0.008557
          ...   
5855    0.010662
5856    0.000137
5857    0.000137
5858   -0.006178
5859    0.008557
Length: 5860, dtype: float64
0       0.004347
1       0.004347
2      -0.004073
3      -0.006178
4       0.008557
          ...   
5855    0.010662
5856    0.000137
5857    0.000137
5858   -0.006178
5859    0.008557
Length: 5860, dtype: float64
0       0.008607
1       0.008607
2      -0.008065
3      -0.012233
4       0.016942
          ...   
5855    0.021110
5856    0.000271
5
```

<!-- grade_id: test5-4_fit, points: 1.3, grade: true, solution: false -->
```python
# 5.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Parts 5. - Test 4 (name: test5-4_fit, points: 1.3)")
print ("\t--->Testing the implementation of 'fit' ...")

file_name_train = 'data' + os.sep + 'car_prices_train.csv'
file_name_test  = 'data' + os.sep + 'car_prices_test.csv'
try:        
    student_name, student_id = myName(), myId()
    assert type(student_name) is str and type(student_id) is int, "name is not a string or id is not an integer"  
    assignment_params = get_assignment_params(student_id)
    X_train, y_train = load_dataset(file_name_train, assignment_params['y_col'], assignment_params['fictive_attr'])
    X_test, y_test   = load_dataset(file_name_test, assignment_params['y_col'], assignment_params['fictive_attr'])
    y_train.name,y_test.name='target','target'
    train_set=pd.concat([X_train,y_train], axis=1)
    test_set =pd.concat([X_test,y_test], axis=1)   
    trained_scaling_info,  train_set_scaled = scale_fit_transform(train_set, assignment_params['fictive_attr'])
    test_set_scaled = scale_transform_for_test(trained_scaling_info,  test_set, assignment_params['fictive_attr'])
    X_train_scaled,y_train_scaled=train_set_scaled.drop(columns=['target']), train_set_scaled['target']
    X_test_scaled,y_test_scaled=test_set_scaled.drop(columns=['target']), test_set_scaled['target']
    alpha,num_iterations = 0.01,10
    trained_w = fit(X_train_scaled,y_train_scaled,alpha,num_iterations)    
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

if 't-distribution standardization' == assignment_params['scale_type']:
    assert -0.014==list(np.round(trained_w.values+0.00005,3))[7], 'Wrong output for weight vector after 10th round'
else:    
    assert 0.05==list(np.round(trained_w.values-0.00055,3))[6], 'Wrong output for weight vector after 10th round'
    
print ("\nGood Job!\nYou've passed the 4th test for the 'fit' function implementation :-)")
```

Output:
```
Parts 5. - Test 4 (name: test5-4_fit, points: 1.3)
	--->Testing the implementation of 'fit' ...
0       0.0
1       0.0
2       0.0
3       0.0
4       0.0
       ... 
5855    0.0
5856    0.0
5857    0.0
5858    0.0
5859    0.0
Length: 5860, dtype: float64
0       0.019462
1       0.029708
2      -0.016163
3      -0.026668
4       0.115535
          ...   
5855   -0.015995
5856    0.012192
5857   -0.028787
5858   -0.025482
5859   -0.007247
Length: 5860, dtype: float64
0       0.038025
1       0.057536
2      -0.031533
3      -0.051945
4       0.225343
          ...   
5855   -0.030921
5856    0.023265
5857   -0.055911
5858   -0.049740
5859   -0.013950
Length: 5860, dtype: float64
0       0.055735
1       0.083595
2      -0.046153
3      -0.075906
4       0.329717
          ...   
5855   -0.044847
5856    0.033297
5857   -0.081470
5858   -0.072836
5859   -0.020146
Length: 5860, dtype: float64
0       0.072638
1       0.107993
2      -0.060061
3      -0.098621
4       0.428935
          
```

<!-- grade_id: test5-5_fit, points: 1.2, grade: true, solution: false -->
```python
# 5.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Parts 5. - Test 5 (name: test5-5_fit, points: 1.2)")
print ("\t--->Testing the implementation of 'fit' ...")
print ("\n\t====> Full grading test - the following test is invisible. You will not know if you passed it, before grades are returned to the students")
```

Output:
```
Parts 5. - Test 5 (name: test5-5_fit, points: 1.2)
	--->Testing the implementation of 'fit' ...

	====> Full grading test - the following test is invisible. You will not know if you passed it, before grades are returned to the students
```

[Go to the beginning of the assignment](#home)
<a id="evaluation"></a>
### 6. Evaluation
This section includes 2 function implementations:
* 6.a. SAE, MAE or RMSE (depending on your assigned method ahead)
* 6.b. R^2 (for everyone)

### Note: all the following <u>metrics are explained in</u> the presentations of the <u>9th lecture and 9th  exercise</u>
This includes all the following metrics:<br />
SAE, MAE, SSE, MSE, RMSE, SST or R^2

<b>Regression - Evaluation (6.a., 6.b.) - Allowed method usage</b>:<br />
<div class="alert alert-success">
&#x2714; Any other basic vector or scalar mathematical function<br />
&#x2714; Any statistical function<br />
</div>
<hr>
<b>Regression - Evaluation (6.a., 6.b.) - Illegal method usage</b>:<br/>
<div class="alert alert-danger">
&#x2718; Any function, which calculates SAE, MAE, SSE, MSE, RMSE, SST or R^2
</div>

### Part 6.a. Instructions
<u>Part 6.a.</u> - Evaluation regression;
<u>method name</u>: <b>evaluate_regression</b><br/>

<pre>
Complete the following function to return the calculated performance evaluation value, 
from the evaluation metric, assigned to you ahead (SAE, MAE or RMSE).
       Note: The evaluation is calculated between y_test (the actual test values) and y_predicted (the predicted values).
------------
input parameter:
- y_test - a series of containing all actual target values per test instance
- y_predicted -  a series of containing all predicted target values per test instance
------------
return value:
- evaluate_value - the (float) value calculated by the evaluation metric (assigned for you ahead)
------------
The return statement should look like the following:
return evaluation_value
</pre>

#### Possible input 
<u>y_test series</u>:

| | |
| :- | :- |
| 0 | 2.271517751  |
| 1 | -0.774988409  |
| 2 | -1.496529342  |
| ... | ...  |


<u>y_predicted</u>:

| | |
| :- | :- |
| 0 | -0.617021277 |
| 1 | -1.202568221 |
| 2 | -0.935330839 |
| ... | ...  |


#### Possible output
1.2421

### 6.a. Regression Evaluation (6.a.) - you need to <u>implement the following method</u> ...

```python
## 6.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# === CELL TYPE: IMPLEMENTATION TASKS

try:    
    student_name, student_id = myName(), myId()
    assert type(student_name) is str and type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)
except Exception as e:
    print ('You probably did not implement student-info functions, \nerror Message:',str(e))
    raise
# ---------------------------
                    
print ("The evaluation metric you need to implement is: " + assignment_params['eval_metric'])
```

Output:
```
The evaluation metric you need to implement is: RMSE
```

<!-- grade_id: cell-e2c529031b105bdd, points: 0, grade: false, solution: true -->
```python
# 6.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ASSISTANCE TO ANSWER 
# ---- Add assistance code here IF NEEDED:
# YOUR CODE HERE

```

<!-- grade_id: answer-evaluate_regression, points: 0, grade: false, solution: true -->
```python
# 6.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ANSWER 

def evaluate_regression(y_test, y_predicted):
    sum1 = 0
    for i in range(y_test.shape[0]):
        sum1 += (y_test.iloc[i] - y_predicted.iloc[i])**2
    sum1 = sum1 * (1/y_test.shape[0])
    evaluation_value = math.sqrt(sum1)
    return evaluation_value
```

```python
# 6.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: SELF TESTING
# Use the following code to test your implementation:
# ---- load test-set:
file_name_test  = 'data' + os.sep + 'car_prices_test.csv'
student_name, student_id = myName(), myId()
assignment_params = get_assignment_params(student_id)
X_test, y_test   = load_dataset(file_name_test, assignment_params['y_col'], assignment_params['fictive_attr'])
y_partial_test = y_test.iloc[1:6]
y_fake_predicted = pd.Series([100000, 200000, 900000, 300000, 400000],y_partial_test.index)
# ~~~~ CALL YOUR 'evaluate_regression' IMPLEMENTATION: ~~~~ #
evaluation_value = evaluate_regression(y_partial_test,y_fake_predicted)
# --- add additional code to check your code if needed:
# YOUR CODE HERE

```

```python
# 6.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation, if used)
# === CODE TYPE: SELF TESTING
# ---- Add your additional tests here if needed:
# YOUR CODE HERE

```

<!-- grade_id: test6a-1_evaluate_regression, points: 0.2, grade: true, solution: false -->
```python
# 6.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Parts 6.a. - Test 1 (name: test6a-1_evaluate_regression, points: 0.2) - Sanity")
print ("\t--->Testing the implementation of 'evaluate_regression' ...")

file_name_train = 'data' + os.sep + 'car_prices_train.csv'
file_name_test  = 'data' + os.sep + 'car_prices_test.csv'
try:        
    student_name, student_id = myName(), myId()
    assert type(student_name) is str and type(student_id) is int, "name is not a string or id is not an integer"  
    assignment_params = get_assignment_params(student_id)
    X_train, y_train = load_dataset(file_name_train, assignment_params['y_col'], assignment_params['fictive_attr'])
    X_test, y_test   = load_dataset(file_name_test, assignment_params['y_col'], assignment_params['fictive_attr'])
    y_test = y_test.iloc[1:6]
    y_predicted = pd.Series([100000, 200000, 900000, 300000, 400000],y_test.index)

    print ("\t\tTesing a '" + assignment_params['eval_metric'] + "' evaluation metric ...")
    evaluation_value = evaluate_regression(y_test,y_predicted)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

print ("\nGood Job!\nYou've passed the 1st test for the 'evaluate_regression' function implementation :-)")
```

Output:
```
Parts 6.a. - Test 1 (name: test6a-1_evaluate_regression, points: 0.2) - Sanity
	--->Testing the implementation of 'evaluate_regression' ...
		Tesing a 'RMSE' evaluation metric ...

Good Job!
You've passed the 1st test for the 'evaluate_regression' function implementation :-)
```

<!-- grade_id: test6a-2_evaluate_regression, points: 0.5, grade: true, solution: false -->
```python
# 6.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Parts 6.a. - Test 2 (name: test6a-2_evaluate_regression, points: 0.5)")
print ("\t--->Testing the implementation of 'evaluate_regression' ...")

file_name_train = 'data' + os.sep + 'car_prices_train.csv'
file_name_test  = 'data' + os.sep + 'car_prices_test.csv'
try:        
    student_name, student_id = myName(), myId()
    assert type(student_name) is str and type(student_id) is int, "name is not a string or id is not an integer"  
    assignment_params = get_assignment_params(student_id)
    X_train, y_train = load_dataset(file_name_train, assignment_params['y_col'], assignment_params['fictive_attr'])
    X_test, y_test   = load_dataset(file_name_test, assignment_params['y_col'], assignment_params['fictive_attr'])
    y_test = y_test.iloc[1:6]
    y_predicted = pd.Series([100000, 200000, 900000, 300000, 400000],y_test.index)
    print ("\t\tTesing a '" + assignment_params['eval_metric'] + "' evaluation metric ...")
    evaluation_value = evaluate_regression(y_test,y_predicted)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
if assignment_params['eval_metric'] == 'SAE':
    assert 2356==evaluation_value//1000, 'Wrong SAE evaluation value'
elif assignment_params['eval_metric'] == 'MAE':
    assert 4712==evaluation_value//100, 'Wrong MAE evaluation value'
else: # if assignment_params['eval_metric'] == 'RMSE':
    assert 564775==np.round(evaluation_value-0.3,0), 'Wrong RMSE evaluation value'
    
print ("\nGood Job!\nYou've passed the 2nd test for the 'evaluate_regression' function implementation :-)")
```

Output:
```
Parts 6.a. - Test 2 (name: test6a-2_evaluate_regression, points: 0.5)
	--->Testing the implementation of 'evaluate_regression' ...
		Tesing a 'RMSE' evaluation metric ...

Good Job!
You've passed the 2nd test for the 'evaluate_regression' function implementation :-)
```

### Part 6.b. Instructions
<u>Part 6.b.</u> - Calculate the r^2 evaluation performance metric; 
<u>method name</u>: <b>calc_r_squared</b><br/>

<pre>
Complete the following function to return the calculated r^2 evaluation value, 
from the r^2 evaluation metric, assigned to you ahead (SAE, MAE or RMSE).
       Note: The evaluation is calculated between y_test (the actual test values) and y_predicted (the predicted values).
------------
input parameter:
- y_test - a series of containing all actual target values per test instance
- y_predicted -  a series of containing all predicted target values per test instance
------------
return value:
- evaluation_value - the r^2 value you've calculated (float number)
------------
The return statement should look like the following:
return evaluate_value
</pre>

[Go to the beginning of the assignment](#home)
<a id="calc_r_square"></a>
### 6.b. The R-squared (R^2) regression evaluation metric

#### Possible input 
<u>y_test series</u>:

| | |
| :- | :- |
| 0 | 2.271517751  |
| 1 | -0.774988409  |
| 2 | -1.496529342  |
| ... | ...  |


<u>y_predicted</u>:

| | |
| :- | :- |
| 0 | -0.617021277 |
| 1 | -1.202568221 |
| 2 | -0.935330839 |
| ... | ...  |


#### Possible output
0.345345

<!-- grade_id: cell-95fb09e1930a31f5, points: 0, grade: false, solution: true -->
```python
# 6.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ASSISTANCE TO ANSWER 
# ---- Add assistance code here IF NEEDED:
# YOUR CODE HERE

```

<!-- grade_id: answer-calc_r_squared, points: 0, grade: false, solution: true -->
```python
# 6.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ANSWER 

def calc_r_squared(y_test,y_predicted):
    sse = 0.0
    yavg = 0.0
    sst = 0.0
    for i in range(y_test.shape[0]):
        sse += (y_test.iloc[i] - y_predicted.iloc[i])**2
        yavg += y_test.iloc[i]
    yavg = yavg * (1 / y_test.shape[0])
    for i in range(y_test.shape[0]):
        sst += (y_test.iloc[i] - yavg)**2
    evaluate_value = 1-( sse / sst)  
    return evaluate_value
```

```python
# 6.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: SELF TESTING
# Use the following code to test your implementation:
# ---- load test-set:
file_name_test  = 'data' + os.sep + 'car_prices_test.csv'
student_name, student_id = myName(), myId()
assignment_params = get_assignment_params(student_id)
X_test, y_test   = load_dataset(file_name_test, assignment_params['y_col'], assignment_params['fictive_attr'])
# ---- fake predict:
y_fake_test = y_test.iloc[2:7]/10000
y_fake_predicted = pd.Series([50, 10, 40, 30, 40],y_fake_test.index)
# ---- evaluate fake predict:
# ~~~~ CALL YOUR 'calc_r_squared' IMPLEMENTATION: ~~~~ #
evaluation_value = calc_r_squared(y_fake_predicted,y_fake_predicted)
# --- add additional code to check your code if needed:
# YOUR CODE HERE

```

```python
# 6.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation, if used)
# === CODE TYPE: SELF TESTING
# ---- Add your additional tests here if needed:
# YOUR CODE HERE

```

<!-- grade_id: test6b-1_calc_r_squared, points: 0.1, grade: true, solution: false -->
```python
# 6.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Parts 6.b. - Test 2 (name: test6b-1_calc_r_squared, points: 0.1) - Sanity")
print ("\t--->Testing the implementation of 'calc_r_squared' ...")

file_name_train = 'data' + os.sep + 'car_prices_train.csv'
file_name_test  = 'data' + os.sep + 'car_prices_test.csv'
try:        
    student_name, student_id = myName(), myId()
    assert type(student_name) is str and type(student_id) is int, "name is not a string or id is not an integer"  
    assignment_params = get_assignment_params(student_id)
    X_train, y_train = load_dataset(file_name_train, assignment_params['y_col'], assignment_params['fictive_attr'])
    X_test, y_test   = load_dataset(file_name_test, assignment_params['y_col'], assignment_params['fictive_attr'])
    y_test = y_test.iloc[2:7]/10000
    y_predicted = pd.Series([50, 10, 40, 30, 40],y_test.index)
    evaluation_value = calc_r_squared(y_test,y_predicted)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

print ("\nGood Job!\nYou've passed the 1st test for the 'calc_r_squared' function implementation :-)")
```

Output:
```
Parts 6.b. - Test 2 (name: test6b-1_calc_r_squared, points: 0.1) - Sanity
	--->Testing the implementation of 'calc_r_squared' ...

Good Job!
You've passed the 1st test for the 'calc_r_squared' function implementation :-)
```

<!-- grade_id: test6b-2_calc_r_squared, points: 0.4, grade: true, solution: false -->
```python
# 6.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Parts 6.b. - Test 2 (name: test6b-2_calc_r_squared, points: 0.4)")
print ("\t--->Testing the implementation of 'calc_r_squared' ...")

file_name_train = 'data' + os.sep + 'car_prices_train.csv'
file_name_test  = 'data' + os.sep + 'car_prices_test.csv'
try:        
    student_name, student_id = myName(), myId()
    assert type(student_name) is str and type(student_id) is int, "name is not a string or id is not an integer"  
    assignment_params = get_assignment_params(student_id)
    X_train, y_train = load_dataset(file_name_train, assignment_params['y_col'], assignment_params['fictive_attr'])
    X_test, y_test   = load_dataset(file_name_test, assignment_params['y_col'], assignment_params['fictive_attr'])
    y_test = y_test.iloc[2:7]/10000
    y_predicted = pd.Series([50, 10, 40, 30, 40],y_test.index)
    evaluation_value = calc_r_squared(y_test,y_predicted)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
assert 0.158==np.round(evaluation_value-0.0005,3), 'Wrong R^2 evaluation value'

print ("\nGood Job!\nYou've passed the 2nd test for the 'calc_r_squared' function implementation :-)")
```

Output:
```
Parts 6.b. - Test 2 (name: test6b-2_calc_r_squared, points: 0.4)
	--->Testing the implementation of 'calc_r_squared' ...

Good Job!
You've passed the 2nd test for the 'calc_r_squared' function implementation :-)
```

<!-- grade_id: test6b-3_calc_r_squared, points: 0.8, grade: true, solution: false -->
```python
# 6.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Parts 6.b. - Test 3 (name: test6b-3_calc_r_squared, points: 0.8)")
print ("\t--->Testing the implementation of 'calc_r_squared' ...")

file_name_train = 'data' + os.sep + 'car_prices_train.csv'
file_name_test  = 'data' + os.sep + 'car_prices_test.csv'
try:        
    student_name, student_id = myName(), myId()
    assert type(student_name) is str and type(student_id) is int, "name is not a string or id is not an integer"  
    assignment_params = get_assignment_params(student_id)
    X_train, y_train = load_dataset(file_name_train, assignment_params['y_col'], assignment_params['fictive_attr'])
    X_test, y_test   = load_dataset(file_name_test, assignment_params['y_col'], assignment_params['fictive_attr'])
    y_train.name,y_test.name='target','target'
    train_set=pd.concat([X_train,y_train], axis=1)
    test_set =pd.concat([X_test,y_test], axis=1)   
    trained_scaling_info,  train_set_scaled = scale_fit_transform(train_set, assignment_params['fictive_attr'])
    test_set_scaled = scale_transform_for_test(trained_scaling_info,  test_set, assignment_params['fictive_attr'])
    X_train_scaled,y_train_scaled=train_set_scaled.drop(columns=['target']), train_set_scaled['target']
    X_test_scaled,y_test_scaled=test_set_scaled.drop(columns=['target']), test_set_scaled['target']
    alpha,num_iterations = 0.01,100
    trained_w = fit(X_train_scaled,y_train_scaled,alpha,num_iterations)
    y_predicted = predict(X_test_scaled, trained_w)
    evaluation_value = calc_r_squared(y_test_scaled,y_predicted)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

print ("\t\tTesing a '" + assignment_params['scale_type'] + "' scale type ...")
if 't-distribution standardization' == assignment_params['scale_type']:
    assert 0.689==np.round(evaluation_value-0.0002,3), 'Wrong R^2 evaluation value'
else: # if 'minmax normalization' == assignment_params['scale_type']:
    assert 0.338==np.round(evaluation_value+0.0001,3), 'Wrong R^2 evaluation value'

print ("\nGood Job!\nYou've passed the 3rd test for the 'calc_r_squared' function implementation :-)")
```

Output:
```
Parts 6.b. - Test 3 (name: test6b-3_calc_r_squared, points: 0.8)
	--->Testing the implementation of 'calc_r_squared' ...
		Tesing a 't-distribution standardization' scale type ...

Good Job!
You've passed the 3rd test for the 'calc_r_squared' function implementation :-)
```

```python

```